```python
# Initialize Otter
import otter
grader = otter.Notebook("hw4.ipynb")
```

---

<h1><center>SDSE Homework 4<br><br> Simple linear regression </center></h1>

---

```python
import numpy as np
import scipy.stats as stats
import matplotlib.pyplot as plt
import pickle
from resources.hashutils import *
```

# 0. Load the data

In this homework we will implement simple linear regression and apply it to four different datasets. 

All of the formulas can be found in the lecture and reader. 

The four datasets are contained in `hw4data.pickle`. Each dataset is a tuple with entries `x` and `y`:  `DA = (x,y)`.

```python
with open('resources/hw4data.pickle','rb') as f:
    DA, DB, DC, DD = pickle.load(f)
```

# 1. Scatter plots

Display each of the four datasets using matplotlib [scatter](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.scatter.html) plots. Your plot should look similar to the one below, with two rows and two columns of plots.  

<img src="resources/f1.png" />

```python
fig, ax = plt.subplots(nrows=...,ncols=...)  # TODO
ax[...,...].scatter(...,...,marker='.')
...
```

# 2. Sample correlation coefficient

Implement the function `corr(D)` that takes a dataset tuple `D` and returns its sample correlation coefficient $r$. 

```python
def corr(D):
    x, y = D
    ...             # TODO
    return ...      # TODO
```

```python
rA = corr(DA)
rB = corr(DB)
rC = corr(DC)
rD = corr(DD)

print(f'rA={rA}')
print(f'rB={rB}')
print(f'rC={rC}')
print(f'rD={rD}')
```

```python
grader.check("q2")
```

# 3. Compute linear regresssion coefficients

Implement the function `compute_linear_regression(D)` that takes a dataset tuple `D` and returns the coefficients of the linear regression $\theta_0$ and $\theta_1$. 

```python
def compute_linear_regression(D):
    ...                 # TODO
    thetahat1 = ...        # TODO
    thetahat0 = ...        # TODO
    return thetahat0, thetahat1
```

```python
thetahatA = compute_linear_regression(DA)
thetahatB = compute_linear_regression(DB)
thetahatC = compute_linear_regression(DC)
thetahatD = compute_linear_regression(DD)

print(f'A: {thetahatA}')
print(f'B: {thetahatB}')
print(f'C: {thetahatC}')
print(f'D: {thetahatD}')
```

```python
grader.check("q3")
```

# 4. Evaluate linear regresssion prediction

Implement the function `eval_linreg(theta,D)` that takes parameters $\hat\theta=(\hat\theta_0,\hat\theta_1)$ and the dataset `D=(x,y)` and returns the array of predictions $\hat{y}$ for the training data: 
$$\hat{y} = \hat\theta_0 + x \hat\theta_1$$

```python
def eval_linreg(thetahat,D):
    ...             # TODO
    return ...      # TODO
```

```python
yhatA = eval_linreg(thetahatA,DA)
yhatB = eval_linreg(thetahatB,DB)
yhatC = eval_linreg(thetahatC,DC)
yhatD = eval_linreg(thetahatD,DD)

print(f'yhatA[:20:3]={yhatA[:20:3]}')
print(f'yhatB[:20:3]={yhatB[:20:3]}')
print(f'yhatC[:20:3]={yhatC[:20:3]}')
print(f'yhatD[:20:3]={yhatD[:20:3]}')
```

```python
grader.check("q4")
```

# 5. Plot linear regresssion prediction

Overlay the linear regression lines on each one of the scatter plots. Your plot should look like this:

<img src="resources/f2.png" />

```python
fig, ax = plt.subplots(figsize=(12,8),nrows=2,ncols=2)
ax[0,0].scatter(...,...,color='k',marker='.')
ax[0,0].plot(...,...)
...
```

# 6. Coefficient of determinination

Implement the function `eval_Rsquared(theta,D)` that takes parameters $\hat\theta=(\hat\theta_0,\hat\theta_1)$ and the dataset `D` and returns the coefficient of determination $R^2$. Verify the identity $R^2=r^2$


```python
def eval_Rsquared(thetahat,D):
    ...             # TODO
    return ...      # TODO
```

```python
RsquaredA = eval_Rsquared(thetahatA,DA)
RsquaredB = eval_Rsquared(thetahatB,DB)
RsquaredC = eval_Rsquared(thetahatC,DC)
RsquaredD = eval_Rsquared(thetahatD,DD)

print(f'RsquaredA={RsquaredA}')
print(f'RsquaredB={RsquaredB}')
print(f'RsquaredC={RsquaredC}')
print(f'RsquaredD={RsquaredD}')
```

```python
grader.check("q6")
```

# 7. Estimate $\sigma^2$

Implement the function `estimate_sigma2(theta,D)` that takes parameters $\hat\theta=(\hat\theta_0,\hat\theta_1)$ and the dataset `D` and returns an unbiased estimate of the variance $\sigma^2$ of the uncertainty values $\varepsilon_i$ in $y_i=\theta_0 + x_i\theta_1  + \varepsilon_i$. 

$$\hat\sigma^2  = \frac{1}{N-2} \sum_{i=1}^{N} (y_i - \hat{y}_i)^2$$

```python
def estimate_sigma2(thetahat,D):
    ...             # TODO
    return ...      # TODO
```

```python
sigma2hatA = estimate_sigma2(thetahatA,DA)
sigma2hatB = estimate_sigma2(thetahatB,DB)
sigma2hatC = estimate_sigma2(thetahatC,DC)
sigma2hatD = estimate_sigma2(thetahatD,DD)

print(f'sigma2A={sigma2hatA}')
print(f'sigma2B={sigma2hatB}')
print(f'sigma2C={sigma2hatC}')
print(f'sigma2D={sigma2hatD}')
```

```python
grader.check("q7")
```

# 8. Uncertainty in the linear regression coefficients

Implement the function `estimate_sigma2(theta,D)` that takes parameters $\hat\theta=(\hat\theta_0,\hat\theta_1)$ and the dataset `D` and returns the variance associated with the linear regression parameter estimates $Var[\hat\Theta_0]$ and $Var[\hat\Theta_1]$.

```python
def compute_var_param(thetahat,D):
    ...             # TODO
    var0 = ...      # TODO
    var1 = ...      # TODO
    return var0, var1
```

```python
varA = compute_var_param(thetahatA,DA)
varB = compute_var_param(thetahatB,DB)
varC = compute_var_param(thetahatC,DC)
varD = compute_var_param(thetahatD,DD)

print(f'A: {varA}')
print(f'B: {varB}')
print(f'C: {varC}')
print(f'D: {varD}')
```

```python
grader.check("q8")
```

# 9. Confidence intervals

Implement the function `compute_rho(theta,D,gamma)` that takes parameters $\hat\theta=(\hat\theta_0,\hat\theta_1)$, the dataset `D`, and a confidence level $\gamma$, and returns the radiuses $\rho_0$ and $\rho_1$ of the $\gamma$-confidence intervals for $\theta_0$ and $\theta_1$ respectively. 

```python
def compute_rho(thetahat,D,gamma):
    ...             # TODO
    rho0 = ...      # TODO
    rho1 = ...      # TODO
    return rho0, rho1
```

```python
gamma = 0.98
rhoA = compute_rho(thetahatA,DA,gamma)
rhoB = compute_rho(thetahatB,DB,gamma)
rhoC = compute_rho(thetahatC,DC,gamma)
rhoD = compute_rho(thetahatD,DD,gamma)

print(f'rhoA={rhoA}')
print(f'rhoB={rhoB}')
print(f'rhoC={rhoC}')
print(f'rhoD={rhoD}')
```

```python
grader.check("q9")
```

# 10. Negligible parameters

We consider a slope parameter to be considered "negligible" when it is close to zero. We define "close" as its 95\% confidence interval including zero. Which of the 4 cases have negligible slope parameter? Provide your answer as a comma-separated string, e.g. `"A,B"`.

```python
def is_negligible(thetahat,D,gamma):
    ...
    captures_zero = ...  # a boolean that is true if 0 is within the confidence interval
    return captures_zero
gamma = 0.95
print(is_negligible(thetahatA,DA,gamma))
print(is_negligible(thetahatB,DB,gamma))
print(is_negligible(thetahatC,DC,gamma))
print(is_negligible(thetahatD,DD,gamma))
answer = ...      # TODO
```

```python
grader.check("q10")
```

# 11. Nonlinear feature vector

The dataset `DA` looks like it could be modeled as an exponential decay. Let's use this insight to improve the regression model for this dataset. Define a feture vector `phi` with
$$\phi(x) = e^{\lambda x}$$
Using trial-and-error, find a value for $\lambda$ that achieves an $R^2$ of 0.9 or more. 

**Note**: The test for this question is hidden and will only be evaluated in Gradescope.

```python
lbda = ...              # TODO
x, y = DA
phi = np.exp(lbda*x)
DA2 = (phi,y)
thetahatA2 = ...           # TODO
RsquaredA2 = ...        # TODO

fig, ax = plt.subplots()
ax.scatter(DA[0],DA[1],color='k',marker='.')
ax.plot(DA[0],eval_linreg(thetahatA,DA),label='$R^2$={:.3f}'.format(RsquaredA))
ax.plot(DA[0],eval_linreg(thetahatA2,DA2),'-',label='$R^2$={:.3f}'.format(RsquaredA2))
ax.legend(fontsize=14)
```

# 12. Prediction uncertainty


Implement the function `compute_var_prediction(theta,D,xi)` that takes parameters $\hat\theta$, the dataset `D`, and a data sample $x_i$, and returns the variance in the prediction $\hat{y}_i$ corresponding to $x_i$.

```python
def compute_var_prediction(thetahat,D,xi):
    ...             # TODO
    return ...      # TODO
```

```python
varYA = compute_var_prediction(thetahatA,DA,DA[0])
stddevYA = np.sqrt(varYA)

varYB = compute_var_prediction(thetahatB,DB,DB[0])
stddevYB = np.sqrt(varYB)

varYC = compute_var_prediction(thetahatC,DC,DC[0])
stddevYC = np.sqrt(varYC)

varYD = compute_var_prediction(thetahatD,DD,DD[0])
stddevYD = np.sqrt(varYD)

print(f'stddevYA={stddevYA[:20:3]}')
print(f'stddevYB={stddevYB[:20:3]}')
print(f'stddevYC={stddevYC[:20:3]}')
print(f'stddevYD={stddevYD[:20:3]}')
```

```python
grader.check("q12")
```

# 13. Plot the prediction with uncertainty bounds

Repeat the scatter and prediction plot, but this time overlay an prediction uncertainty band of plus/minus two standard deviations. Your plot should look something like this:

<img src="resources/f3.png" />


```python
fig, ax = plt.subplots(figsize=(12,8),nrows=2,ncols=2)

a = ax[0,0]
yhatA = eval_linreg(thetahatA,DA)
a.scatter(DA[0],DA[1],color='k',marker='.')
a.plot(DA[0],yhatA)
a.fill_between(DA[0],yhatA-2*np.sqrt(varYA),yhatA+2*np.sqrt(varYA),alpha=0.5)

...
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q10 (points: None)
```
>>> get_hash(answer) == 'f623e75af30e62bbd73d6df5b50bb7b5'
True
```


Tests for q11 (points: None)


Tests for q12 (points: None)
```
>>> thetahatA = compute_linear_regression(DA)
>>> varYA = compute_var_prediction(thetahatA, DA, DA[0])
>>> get_hash(varYA, 4) == 'f4db257cd335e98f201ea47d8bf5966c'
True
```
```
>>> thetahatB = compute_linear_regression(DB)
>>> varYB = compute_var_prediction(thetahatB, DB, DB[0])
>>> get_hash(varYB, 4) == '749dfa6987e7f0afe8125549d5a737c3'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> varYC = compute_var_prediction(thetahatC, DC, DC[0])
>>> get_hash(varYC, 4) == '42abd98bfb3f757280a0b8cd904da0c5'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> varYD = compute_var_prediction(thetahatD, DD, DD[0])
>>> get_hash(varYD, 4) == '503b0a53d842c647d84a456dc4ff9434'
True
```


Tests for q2 (points: None)
```
>>> bool(np.isclose(rA, -0.8099147715792736, 0.001))
True
```
```
>>> get_hash(rB, 4) == '835a2b7b610ff0f3cbc96d7b636415b7'
True
```


Tests for q3 (points: None)
```
>>> thetahatA = compute_linear_regression(DA)
>>> get_hash(thetahatA, 4) == 'ef0f6da7edfcca4e28195f3a4d30a793'
True
```
```
>>> thetahatB = compute_linear_regression(DB)
>>> get_hash(thetahatB, 4) == '53bdcb310916d5082434d92adecf9ec5'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> get_hash(thetahatC, 4) == '34147028d6b5b5da4d98154e73ec4888'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> get_hash(thetahatD, 4) == 'd24429496a8468962395ba9d00e40c20'
True
```


Tests for q4 (points: None)
```
>>> thetahatA = compute_linear_regression(DA)
>>> yhatA = eval_linreg(thetahatA, DA)
>>> get_hash(yhatA, 4) == '85c6f31c18559d1d6a77afb2394cdbc9'
True
```
```
>>> thetahatB = compute_linear_regression(DB)
>>> yhatB = eval_linreg(thetahatB, DB)
>>> get_hash(yhatB, 4) == 'b9a8973df0ca79c46ef2b71b4c702ae2'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> yhatC = eval_linreg(thetahatC, DC)
>>> get_hash(yhatC, 4) == '8bd3f7c0bf9cc0cfbf30f5bdbd78457e'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> yhatD = eval_linreg(thetahatD, DD)
>>> get_hash(yhatD, 4) == 'c2a4783c823ea82124b3db5b4ddbe762'
True
```


Tests for q6 (points: None)
```
>>> thetahatA = compute_linear_regression(DA)
>>> RsquaredA = eval_Rsquared(thetahatA, DA)
>>> get_hash(RsquaredA, 4) == 'a342403c381016c41ecbc12123699daa'
True
```
```
>>> thetahatB = compute_linear_regression(DB)
>>> RsquaredB = eval_Rsquared(thetahatB, DB)
>>> get_hash(RsquaredB, 4) == '9bb31e2afd1462aac7abc171d7533791'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> RsquaredC = eval_Rsquared(thetahatC, DC)
>>> get_hash(RsquaredC, 4) == '9637bf59ecb1b3c0d15f8e1628a083f4'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> RsquaredD = eval_Rsquared(thetahatD, DD)
>>> get_hash(RsquaredD, 4) == '608bbf3ac101b1f6da9ee89f11e8e564'
True
```


Tests for q7 (points: None)
```
>>> thetahatA = compute_linear_regression(DA)
>>> sigma2hatA = estimate_sigma2(thetahatA, DA)
>>> get_hash(sigma2hatA, 4) == '453bf4d9934d0e3874a71cff19ba92bd'
True
```
```
>>> thetahatB = compute_linear_regression(DB)
>>> sigma2hatB = estimate_sigma2(thetahatB, DB)
>>> get_hash(sigma2hatB, 4) == 'd7220b6c2c60895fd32ffff7577b6c95'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> sigma2hatC = estimate_sigma2(thetahatC, DC)
>>> get_hash(sigma2hatC, 4) == '06265845401f290b89a5a6717fe6e546'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> sigma2hatD = estimate_sigma2(thetahatD, DD)
>>> get_hash(sigma2hatD, 4) == '36552b079970ffb2dd1314115af76c4b'
True
```


Tests for q8 (points: None)
```
>>> thetahatA = compute_linear_regression(DA)
>>> varA = compute_var_param(thetahatA, DA)
>>> get_hash(varA, 4) == '8743af846095cf842181add140dc8a82'
True
```
```
>>> thetahatA = compute_linear_regression(DB)
>>> varB = compute_var_param(thetahatB, DB)
>>> get_hash(varB, 4) == '9812d81a6797c0711fae0fe51455a37f'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> varC = compute_var_param(thetahatC, DC)
>>> get_hash(varC, 4) == 'edd0658c956ba0f360a012c107c3de14'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> varD = compute_var_param(thetahatD, DD)
>>> get_hash(varD, 4) == '9cf6ea9abf256bd7144bc5bc753f3d80'
True
```


Tests for q9 (points: None)
```
>>> thetahatA = compute_linear_regression(DA)
>>> rhoA = compute_rho(thetahatA, DA, gamma)
>>> get_hash(rhoA, 4) == '61fb843f8c14617d0e3c7ac9e976d81c'
True
```
```
>>> thetahatB = compute_linear_regression(DB)
>>> rhoB = compute_rho(thetahatB, DB, gamma)
>>> get_hash(rhoB, 4) == '265119261e682e2d5e667810484596bd'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> rhoC = compute_rho(thetahatC, DC, gamma)
>>> get_hash(rhoC, 4) == 'd2ccbeed10d93b3bae5c2b7bcf0fff75'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> rhoD = compute_rho(thetahatD, DD, gamma)
>>> get_hash(rhoD, 4) == 'ebf3a54b7e494cd7b2d856545b58a488'
True
```
