```python
# Initialize Otter
import otter
grader = otter.Notebook("hw4.ipynb")
```

---

<h1><center>SDSE Homework 4<br><br> Simple linear regression </center></h1>

---

```python
import numpy as np
import scipy.stats as stats
import matplotlib.pyplot as plt
import pickle
from resources.hashutils import *
```

# 0. Load the data

In this homework we will implement simple linear regression and apply it to four different datasets. 

All of the formulas can be found in the lecture and reader. 

The four datasets are contained in `resources/hw4data.pickle`. Each dataset is a tuple with entries `x` and `y`:  `DA = (x,y)`.

```python
with open('resources/hw4data.pickle','rb') as f:
    DA, DB, DC, DD = pickle.load(f)
```

```python
# Use type() and .shape to explore the contents of DA through DD

```

# 1. Scatter plots

Display each of the four datasets using matplotlib [scatter](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.scatter.html) plots. Your plot should look similar to the one below, with two rows and two columns of plots.  

<img src="resources/f1.png" />

```python
fig, ax = plt.subplots(nrows=...,ncols=...)
ax[...,...].scatter(...,...,color='k',marker='.')   # plot DA in the upper left corner
...   # plot DB in the upper right corner
...   # plot DC in the lower left corner
...   # plot DD in the lower right corner
```

# 2. Sample correlation coefficient

Implement the function `corr(D)` that takes a dataset tuple `D` and returns its sample correlation coefficient $r$. 

```python
def corr(D):
    ...             # TODO
    return ...      # TODO
```

```python
rA = corr(DA)
rB = corr(DB)
rC = corr(DC)
rD = corr(DD)

print(f'rA = {rA}')
print(f'rB = {rB}')
print(f'rC = {rC}')
print(f'rD = {rD}')
```

```python
grader.check("q2")
```

# 3. Compute linear regresssion coefficients

Implement the function `compute_linear_regression(D)` that takes a dataset tuple `D` and returns the coefficients of the linear regression $\hat{\theta}_0$ and $\hat{\theta}_1$. 

```python
def compute_linear_regression(D):
    ...                 # TODO
    thetahat1 = ...        # TODO
    thetahat0 = ...        # TODO
    return thetahat0, thetahat1
```

```python
thetahat0A, thetahat1A = compute_linear_regression(DA)
thetahat0B, thetahat1B = compute_linear_regression(DB)
thetahat0C, thetahat1C = compute_linear_regression(DC)
thetahat0D, thetahat1D = compute_linear_regression(DD)

print(f'A: thetahat0 = {thetahat0A:.2e}, thetahat1 = {thetahat1A:.2e}')
print(f'B: thetahat0 = {thetahat0B:.2e}, thetahat1 = {thetahat1B:.2e}')
print(f'C: thetahat0 = {thetahat0C:.2e}, thetahat1 = {thetahat1C:.2e}')
print(f'D: thetahat0 = {thetahat0D:.2e}, Vhetahat1 = {thetahat1D:.2e}')
```

```python
grader.check("q3")
```

# 4. Evaluate linear regresssion prediction

Implement the function `predict(theta,D)` that takes parameter tuple $\hat\theta=(\hat\theta_0,\hat\theta_1)$ and the dataset `D` and returns the array of predictions $\hat{y}$ for `D`: 
$$\hat{y} = \hat\theta_0 + x \hat\theta_1$$

```python
def predict(thetahat,D):
    ...             # TODO
    return ...      # TODO
```

```python
yhatA = predict(thetahatA,DA)
yhatB = predict(thetahatB,DB)
yhatC = predict(thetahatC,DC)
yhatD = predict(thetahatD,DD)

print(f'yhatA[:20:3]={yhatA[:20:3]}')
print(f'yhatB[:20:3]={yhatB[:20:3]}')
print(f'yhatC[:20:3]={yhatC[:20:3]}')
print(f'yhatD[:20:3]={yhatD[:20:3]}')
```

```python
grader.check("q4")
```

# 5. Plot linear regresssion prediction

Overlay the linear regression lines on each of the scatter plots from plot 1. Your plot should look like this:

<img src="resources/f2.png" />

```python
fig, ax = plt.subplots(figsize=(12,8),nrows=2,ncols=2)
...
```

# 6. Coefficient of determinination

Implement the function `eval_Rsquared(theta,D)` that takes parameters $\hat\theta=(\hat\theta_0,\hat\theta_1)$ and the dataset `D` and returns the coefficient of determination $R^2$ (formula in the lecture on supervised learning). Verify the identity $R^2=r^2$

```python
def eval_Rsquared(thetahat,D):
    ...             # TODO
    return ...      # TODO
```

```python
RsquaredA = eval_Rsquared(thetahatA,DA)
RsquaredB = eval_Rsquared(thetahatB,DB)
RsquaredC = eval_Rsquared(thetahatC,DC)
RsquaredD = eval_Rsquared(thetahatD,DD)

print( np.isclose( RsquaredA , corr(DA)**2 ))
print( np.isclose( RsquaredB , corr(DB)**2 ))
print( np.isclose( RsquaredC , corr(DC)**2 ))
print( np.isclose( RsquaredD , corr(DD)**2 ))
```

```python
grader.check("q6")
```

# 7. Estimate the variance of the output measurement noise ($\sigma^2$)

Implement the function `estimate_sigma2(D)` that takes a dataset `D` and returns the estimated variance of the output measurement noise $\varepsilon_i$. 

$$\hat\sigma^2  = \frac{1}{N-2} \sum_{i=1}^{N} (y_i - \hat{y}_i)^2$$

where $\hat{y}_i$ are the predictions of linear regression for `D`.

```python
def estimate_sigma2(D):
    ...             # TODO
    return ...      # TODO
```

```python
sigma2hatA = estimate_sigma2(DA)
sigma2hatB = estimate_sigma2(DB)
sigma2hatC = estimate_sigma2(DC)
sigma2hatD = estimate_sigma2(DD)

print(f'sigma2A = {sigma2hatA:.2e}')
print(f'sigma2B = {sigma2hatB:.2e}')
print(f'sigma2C = {sigma2hatC:.2e}')
print(f'sigma2D = {sigma2hatD:.2e}')
```

```python
grader.check("q7")
```

# 8. Uncertainty in the linear regression coefficients

Implement the function `compute_var_param(D)` that takes a dataset `D` and returns the variance associated with the linear regression parameters.

$$Var[\hat\Theta_1] = \frac{\hat\sigma^2}{(N-1)\hat{\sigma}_X^2}$$

$$Var[\hat\Theta_0] = \frac{\hat\sigma^2}{N} + \frac{\hat\sigma^2 \hat{\mu}^2_X  }{(N-1)\hat{\sigma}_X^2}$$

```python
def compute_var_param(D):
    ...             # TODO
    var1 = ...      # TODO
    var0 = ...      # TODO
    return var1, var0
```

```python
var1A, var0A = compute_var_param(DA)
var1B, var0B = compute_var_param(DB)
var1C, var0C = compute_var_param(DC)
var1D, var0D = compute_var_param(DD)

print(f'A: Var[thetahat1] = {var1A:.2e}, Var[thetahat0] = {var0A:.2e}')
print(f'B: Var[thetahat1] = {var1B:.2e}, Var[thetahat0] = {var0B:.2e}')
print(f'C: Var[thetahat1] = {var1C:.2e}, Var[thetahat0] = {var0C:.2e}')
print(f'D: Var[thetahat1] = {var1D:.2e}, Var[thetahat0] = {var0D:.2e}')
```

```python
grader.check("q8")
```

# 9. Confidence intervals

Implement the function `compute_rho(D,gamma)` that takes a dataset `D` and confidence level $\gamma$, and returns the radiuses $\rho_0$ and $\rho_1$ of the $\gamma$-confidence intervals for $\hat\theta_0$ and $\hat\theta_1$ respectively.

Note:
+ See ``Linear Regression part 1'' for the formulas.
+ Because $N$ is ``large'' ($N>40$) we can use the inverse cdf of the unit normal distribution. 

```python
def compute_rho(D,gamma):
    ...             # TODO
    rho1 = ...      # TODO
    rho0 = ...      # TODO
    return rho1, rho0
```

```python
gamma = 0.98
rho1A, rho0A = compute_rho(DA,gamma)
rho1B, rho0B = compute_rho(DB,gamma)
rho1C, rho0C = compute_rho(DC,gamma)
rho1D, rho0D = compute_rho(DD,gamma)

print(f'A: rho1 = {rho1A:.2e}, rho0 = {rho0A:.2e}')
print(f'B: rho1 = {rho1B:.2e}, rho0 = {rho0B:.2e}')
print(f'C: rho1 = {rho1C:.2e}, rho0 = {rho0C:.2e}')
print(f'D: rho1 = {rho1D:.2e}, rho0 = {rho0D:.2e}')
```

```python
grader.check("q9")
```

# 10. Negligible parameters

We consider a slope parameter to be "negligible" when it is close to zero. We define "close" to mean that zero is within its 95\% confidence interval. Which of the 4 cases have a negligible slope parameter $\hat\theta_1$? Provide your answer as a comma-separated string, e.g. `"A,B"`.

```python
def is_negligible(thetahat,D,gamma):
    ...
    captures_zero = ...  # a boolean that is true if 0 is within the confidence interval
    return captures_zero

gamma = 0.95
print(is_negligible(DA,gamma))
print(is_negligible(DB,gamma))
print(is_negligible(DC,gamma))
print(is_negligible(DD,gamma))
answer = ...      # TODO
```

```python
grader.check("q10")
```

# 11. Nonlinear feature vector

The dataset `DA` looks like it could be modeled as exponential decay. Let's use this insight to improve the regression model for this dataset. Define a feture vector `phi` with
$$\phi(x) = e^{\lambda x}$$
Using trial-and-error, find a value for $\lambda$ that achieves an $R^2$ of 0.9 or more. 

**Note**: The test for this question is hidden and will only be evaluated in Gradescope.

```python
lbda = ...              # TODO
x, y = DA
phi = np.exp(lbda*x)
DA2 = (phi,y)
thetahatA2 = compute_linear_regression(DA2)    # TODO
RsquaredA2 = eval_Rsquared(thetahatA2,DA2)     # TODO

fig, ax = plt.subplots()
ax.scatter(DA[0],DA[1],color='k',marker='.')
ax.plot(DA[0],predict(thetahatA,DA),label='$R^2$={:.3f}'.format(RsquaredA))
ax.plot(DA[0],predict(thetahatA2,DA2),'-',label='$R^2$={:.3f}'.format(RsquaredA2))
ax.legend(fontsize=14)
```

# 12. Prediction uncertainty

Implement the function `compute_var_prediction(D,x)` that takes a dataset `D` and an array of test input samples `xtest`, and returns the array of variances in the predictions $\hat{y}$ corresponding to elements in `xtest`.

$$Var[\hat{y}] = \frac{\hat\sigma^2}{N} + \frac{\hat\sigma^2(x-\hat\mu_X)^2}{(N-1)\hat{\sigma}^2_X}  $$

```python
def compute_var_prediction(D,xtest):
    ...             # TODO
    return ...      # TODO
```

```python
# Evaluate the prediction variance on the training input samples. 

varYhatA = compute_var_prediction(DA,DA[0])
stddevYhatA = np.sqrt(varYhatA)

varYhatB = compute_var_prediction(DB,DB[0])
stddevYhatB = np.sqrt(varYhatB)

varYhatC = compute_var_prediction(DC,DC[0])
stddevYhatC = np.sqrt(varYhatC)

varYhatD = compute_var_prediction(DD,DD[0])
stddevYhatD = np.sqrt(varYhatD)

print(f'stddevYA={stddevYhatA[:20:3]}')
print(f'stddevYB={stddevYhatB[:20:3]}')
print(f'stddevYC={stddevYhatC[:20:3]}')
print(f'stddevYD={stddevYhatD[:20:3]}')
```

```python
grader.check("q12")
```

# 13. Plot the prediction with uncertainty bounds

Repeat the scatter and prediction plot, but this time overlay an prediction uncertainty band of plus/minus two standard deviations. Your plot should look something like this:

<img src="resources/f3.png" />


```python
fig, ax = plt.subplots(figsize=(12,8),nrows=2,ncols=2)

a = ax[0,0]
yhatA = predict(thetahatA,DA)
a.scatter(DA[0],DA[1],color='k',marker='.')
a.plot(DA[0],yhatA)
a.fill_between(DA[0],yhatA-2*np.sqrt(varYA),yhatA+2*np.sqrt(varYA),alpha=0.5)

...
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q10 (points: None)
```
>>> get_hash(answer) == 'f623e75af30e62bbd73d6df5b50bb7b5'
True
```


Tests for q11 (points: None)


Tests for q12 (points: None)
```
>>> varYA = compute_var_prediction(DA, DA[0])
>>> get_hash(varYA, 4) == 'f4db257cd335e98f201ea47d8bf5966c'
True
```
```
>>> varYB = compute_var_prediction(DB, DB[0])
>>> get_hash(varYB, 4) == '749dfa6987e7f0afe8125549d5a737c3'
True
```
```
>>> varYC = compute_var_prediction(DC, DC[0])
>>> get_hash(varYC, 4) == '42abd98bfb3f757280a0b8cd904da0c5'
True
```
```
>>> varYD = compute_var_prediction(DD, DD[0])
>>> get_hash(varYD, 4) == '503b0a53d842c647d84a456dc4ff9434'
True
```


Tests for q2 (points: None)
```
>>> bool(np.isclose(rA, -0.8099147715792736, 0.001))
True
```
```
>>> get_hash(rB, 4) == '835a2b7b610ff0f3cbc96d7b636415b7'
True
```


Tests for q3 (points: None)
```
>>> thetahat0A, thetahat1A = compute_linear_regression(DA)
>>> bool(thetahat0A > 3 and thetahat1A < 0)
True
```
```
>>> thetahatA = compute_linear_regression(DA)
>>> get_hash(thetahatA, 4) == 'ef0f6da7edfcca4e28195f3a4d30a793'
True
```
```
>>> thetahatB = compute_linear_regression(DB)
>>> get_hash(thetahatB, 4) == '53bdcb310916d5082434d92adecf9ec5'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> get_hash(thetahatC, 4) == '34147028d6b5b5da4d98154e73ec4888'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> get_hash(thetahatD, 4) == 'd24429496a8468962395ba9d00e40c20'
True
```


Tests for q4 (points: None)
```
>>> thetahatA = compute_linear_regression(DA)
>>> yhatA = predict(thetahatA, DA)
>>> get_hash(yhatA, 4) == '85c6f31c18559d1d6a77afb2394cdbc9'
True
```
```
>>> thetahatB = compute_linear_regression(DB)
>>> yhatB = predict(thetahatB, DB)
>>> get_hash(yhatB, 4) == 'b9a8973df0ca79c46ef2b71b4c702ae2'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> yhatC = predict(thetahatC, DC)
>>> get_hash(yhatC, 4) == '8bd3f7c0bf9cc0cfbf30f5bdbd78457e'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> yhatD = predict(thetahatD, DD)
>>> get_hash(yhatD, 4) == 'c2a4783c823ea82124b3db5b4ddbe762'
True
```


Tests for q6 (points: None)
```
>>> thetahatA = compute_linear_regression(DA)
>>> RsquaredA = eval_Rsquared(thetahatA, DA)
>>> get_hash(RsquaredA, 4) == 'a342403c381016c41ecbc12123699daa'
True
```
```
>>> thetahatB = compute_linear_regression(DB)
>>> RsquaredB = eval_Rsquared(thetahatB, DB)
>>> get_hash(RsquaredB, 4) == '9bb31e2afd1462aac7abc171d7533791'
True
```
```
>>> thetahatC = compute_linear_regression(DC)
>>> RsquaredC = eval_Rsquared(thetahatC, DC)
>>> get_hash(RsquaredC, 4) == '9637bf59ecb1b3c0d15f8e1628a083f4'
True
```
```
>>> thetahatD = compute_linear_regression(DD)
>>> RsquaredD = eval_Rsquared(thetahatD, DD)
>>> get_hash(RsquaredD, 4) == '608bbf3ac101b1f6da9ee89f11e8e564'
True
```


Tests for q7 (points: None)
```
>>> sigma2hatA = estimate_sigma2(DA)
>>> get_hash(sigma2hatA, 4) == '453bf4d9934d0e3874a71cff19ba92bd'
True
```
```
>>> sigma2hatB = estimate_sigma2(DB)
>>> get_hash(sigma2hatB, 4) == 'd7220b6c2c60895fd32ffff7577b6c95'
True
```
```
>>> sigma2hatC = estimate_sigma2(DC)
>>> get_hash(sigma2hatC, 4) == '06265845401f290b89a5a6717fe6e546'
True
```
```
>>> sigma2hatD = estimate_sigma2(DD)
>>> get_hash(sigma2hatD, 4) == '36552b079970ffb2dd1314115af76c4b'
True
```


Tests for q8 (points: None)
```
>>> var1A, var0A = compute_var_param(DA)
>>> bool(var1A < 2e-05 and var0A > 0.0045)
True
```
```
>>> var1A, var0A = compute_var_param(DA)
>>> get_hash((var1A, var0A)) == '3bf247d413aaa60a631f0fc3cfccd009'
True
```
```
>>> var1B, var0B = compute_var_param(DB)
>>> get_hash((var1B, var0B)) == 'd3c76c6d125b164cd9948ccfde13222d'
True
```
```
>>> var1C, var0C = compute_var_param(DC)
>>> get_hash((var1C, var0C)) == '0db4e954045737ebe09813fb6ac93ee5'
True
```
```
>>> var1D, var0D = compute_var_param(DD)
>>> get_hash((var1D, var0D)) == 'a8897718d09f20cbec39a20a1564159f'
True
```


Tests for q9 (points: None)
```
>>> rho1A, rho0A = compute_rho(DA, gamma)
>>> get_hash(rho1A, rho0A) == 'b76d39609922ebc6e0894a30b5550e51'
True
```
```
>>> rho1B, rho0B = compute_rho(DB, gamma)
>>> get_hash(rho1B, rho0B) == '30565a8911a6bb487e3745c0ea3c8224'
True
```
```
>>> rho1C, rho0C = compute_rho(DC, gamma)
>>> get_hash(rho1C, rho0C) == 'f7489f887cd57651d699cb76def574ad'
True
```
```
>>> rho1D, rho0D = compute_rho(DD, gamma)
>>> get_hash(rho1D, rho0D) == 'a5e5db9d32625694fb10c94ac99922ec'
True
```
