```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

---

<h2><center>Demo for Jupyter Notebooks based exercises <br></center></h2>

---

```python
from resources.hashutils import *
```

# Part 1: A simple calculation
Compute the area of a square with sides of length L=2 meters

```python
L = 2.1
area = ...
```

```python
grader.check("q1")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> area > 0 and area < 10
True
```
```
>>> get_hash(area, 4) == '72cd62a3605c1a89ce3ae9260615462b'
True
```
