```python
# Initialize Otter
import otter
grader = otter.Notebook("lab4.ipynb")
```

---

<h1><center>SDSE Lab 4 <br><br> Linear regression and Feature selection </center></h1>

---

In this lab we will use linear regression to predict cancer mortality rates based on data obtained from the American Community Survey of the [U.S. Census Bureau](https://www.census.gov/). The lab has four parts. In part 1 you will load the data and do basic manipulations using [pandas](https://pandas.pydata.org/docs/index.html). Pandas is a Python package that specializes in tabular data. It is widely used in data science and machine learning since the data in these fields are usually structured as a table. Pandas is a very powerful library that is well worth investing some time to learn [here](https://pandas.pydata.org/docs/getting_started/index.html#getting-started) and [here](https://pandas.pydata.org/docs/user_guide/index.html).

In part 2 you will perform linear regression on the full feature set. In part 3 you will compute confidence intervals and perform hypothesis tests on the paramters from part 2. Finally, in part 4 you will run the forward and backward stepwise feature selection algorithms and estimate the performance of the resulting model using a test dataset.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
import scipy.stats as stats
from resources.hashutils import *
```

# Part 1:  Loading and cleaning the data

## 1.1 Load the data into a pandas DataFrame

See a description of the data [here](https://data.world/nrippner/ols-regression-challenge).

Use [pd.read_csv](https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html) to load the data from `resources/cancerdata.csv`.

You can obtain information about the data using these DataFrame methods:
+ `data0.head()` to see the first 5 rows of the DataFrame.
+ `data0.tail()` to see the last 5 rows of the DataFrame.
+ `data0.shape` is the tuple (number of rows, number of columns).
+ `data0.columns` contains the column labels.
+ `data0.index` contains the row labels.

```python
data0 = ...  # TODO
```

## 1.2 Inspect columns

Run `data0.info()` and note:
 a) which inputs are non-numerical (Dtype=object), and
 b) which inputs have null entries (Non-Null Count<3047).

Store the names ('Column' entry) of the non-numerical inputs in the set `non_numerical_inputs`. Store the names of inputs with null entries in the set `null_entry_inputs`.

**Note**: If not all of the rows of `data0.info()` are displayed, you'll probably have this message at the bottom:

*``Output is truncated. View as a scrollable element or open in a text editor. Adjust cell output settings...''*

```python
data0.info()
non_numerical_inputs = {...,...}      # TODO
null_entry_inputs = {...,...,...}     # TODO
```

```python
grader.check("q1p2")
```

## 1.3 Discard non-numerical columns

Remove the two columns with non-numeric data.

Hints:
+ `data0.dtypes` lists the data types for each column.
+ You can construct a boolean indicator array of non-numeric colums with `data0.dtypes=='object'`.
+ Use that array to index `data.columns`
+ Use [`data0.drop`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.drop.html) to remove the selected columns.
+ Save the result as `data1`

```python
ind = ...               # TODO
drop_cols = ...         # TODO
data1 = data0.drop(...)       # TODO
```

```python
grader.check("q1p3")
```

## 1.4 Discard columns where more than 10% of values are nan

Hints:
+ [`.dropna`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.dropna.html)
+ `axis=1` ... we are dropping columns
+ `thresh` ... The condition for dropping a column is that it has less than `round(0.9*data1.shape[0])` non-nans.
+ Save the result as `data2`


```python
thresh = ...                # TODO
data2 = data1.dropna(...)     # TODO
data2.shape
```

```python
grader.check("q1p4")
```

## 1.5 Drop all rows that contain one or more nans.

Save the result as `data3`.

Hint: You can again use `dropna` for this.

```python
data3 = ...      # TODO
```

```python
grader.check("q1p5")
```

## 1.6 Inspect correlations

Next we'll look at the correlation coefficients between each of the inputs and the target variable `target_deathrate`. This is a quick way to check which of the inputs may be most useful to include in the model. Correlations only provide an initial guess, however. Remember that the correlation coefficient only measures the linear relationship between variables. That's perfect when the model is linear (as in this lab activity), but less useful for nonlinear models.

1) Use `data3.corr()` to build the correlations matrix.
2) Inspect the column (or row) corresponding to `target_deathrate`.
3) Rank (i.e. sort) the inputs from most to least correlated with the output. This ranking is in terms of the absolute value of the correlation coefficient.
4) Save the top 5 correlated inputs to `top_5_corr`. `top_5_corr` should be a numpy array with shape `(5,)`.

Hints:
+ [`abs`](https://pandas.pydata.org/docs/reference/api/pandas.Series.abs.html)
+ [`sort_values`](https://pandas.pydata.org/docs/reference/api/pandas.Series.sort_values.html)
+ [`to_numpy`](https://pandas.pydata.org/docs/reference/api/pandas.Series.to_numpy.html)

```python
# correlation matrix
C = ...         # TODO

# vector correlations between the inputs and target_deathrate
corr_target = ...         # TODO

# sorted corr_target_sort
corr_target_sort = ...         # TODO

# top 5 correlations with target_deathrate
top_5_sort = ...         # TODO
```

```python
grader.check("q1p6")
```

## 1.7 Scatter plot

Make a scatter plot of the data with the most correlated input along the x axis, and the target along the y axis.

Hint: You can use the plotting function attached to the DataFrame: [data.plot(kind='scatter',x=..., y=...)](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.plot.html)

```python
data3.plot(kind='scatter',x=...,y=...)   # TODO
```

---

# Part 2: Linear regression 

## 2.1 Extract `X` and `Y` from `data3`

The next cell extracts the `X` and `Y` matrices from `data3`. It also makes a list of `inputs`, and defines the number of samples `N` and the number of inputs `P`.

```python
X = data3.drop(columns='target_deathrate').values
inputs = data3.columns.values
inputs = inputs[inputs!='target_deathrate']
Y = data3['target_deathrate'].values

N = ...
D = ...
```

```python
grader.check("q2p1")
```

## 2.2 Center the inputs

\begin{align*}
\hat\mu_X &= \frac{1}{N} \mathbf{1}^T_N \mathbf{X} \\
\mathbf{X}^c&=\mathbf{X}-\mathbf{1}_N\hat\mu_X
\end{align*}

Compute the column-wise means `muhatX` and subtract them from `X` to obtain `Xc`. 

Hint: 
+ Use `X.mean(axis=...)`. Should it be `axis=0` or `axis=1`? 
+ The formula for `Xc` above has $\mathbf{1}_N\hat\mu_X$. The broadcasting rules of numpy make multiplying $\hat\mu_X$ by $\mathbf{1}_N$ unnecessary.

Check that the column-wise means of `Xc` equal zero (to machine precision). 

```python
muhatX = ...    # TODO
Xc = ...        # TODO
```

```python
grader.check("q2p2")
```

## 2.3 Center the outputs

\begin{align*}
\hat\mu_Y &= \frac{1}{N} \mathbf{1}^T_N \mathbf{Y} \\
\mathbf{Y}^c&=\mathbf{Y}-\mathbf{1}_N\hat\mu_Y
\end{align*}


```python
muhatY = ...  # TODO
Yc = ... # TODO
```

```python
grader.check("q2p3")
```

## 2.4 Compute the inverse input covariance martrix.


$$\Sigma^{-1}_X=( (\mathbf{X}^c)^T \mathbf{X}^c)^{-1}$$

Hints:
+ numpy's @ operator and matrix transposition (`Xc.T`). See [this](https://numpy.org/devdocs/user/numpy-for-matlab-users.html) summary of the key differences between numpy and Matlab.
+ [np.linalg.inv](https://numpy.org/doc/stable/reference/generated/numpy.linalg.inv.html)


```python
invCovX = ...   # TODO
```

```python
grader.check("q2p4")
```

## 2.5 Compute the least squares estimates of the parameters

\begin{align*}
\underline{\hat\theta}_1 &= \Sigma^{-1}_X(\mathbf{X}^c)^{T} \mathbf{Y}^c \\
\hat\theta_0 &= \hat\mu_Y - \hat\mu_X \underline{\hat\theta}_1
\end{align*}

```python
theta1hat = ...  # TODO
theta0hat = ...  # TODO
```

```python
grader.check("q2p5")
```

## 2.6 Compute predictions for each of the samples

$$\mathbf{\hat{Y}} = \mathbf{1}_N\hat\theta_0 + \mathbf{X}\underline{\hat\theta}_1$$


```python
Yhat = ...  # TODO
```

```python
grader.check("q2p6")
```

## 2.7 Model performance

Compute the coefficient of determination $R^2$ for this model on the trainging data.

$$ R^2 = 1 - \frac{\sum_{i=1}^N (y_i-\hat{y}_i)^2}{\sum_{i=1}^N (y_i-\hat\mu_Y)^2}$$

```python
R2 = ...    # TODO
```

```python
grader.check("q2p7")
```

---

# Part 3: Parameter uncertainty

## 3.1 Estimate the variance of the output

$$\hat\sigma^2 = \frac{1}{N-D-1} \sum_{i=1}^{N}(y_i-\hat{y}_i)^2$$

```python
sigmahat2 = ...
```

```python
grader.check("q3p1")
```

## 3.2 Compute the variances of the slope parameters

These are the diagonal entries of 
$$Var[\underline{\hat\Theta}_1] = \hat\sigma^2 \Sigma^{-1}_X $$

In other words, the variance of $\hat\theta_p$ is the $p$'th diagonal entry of $Var[\underline{\hat\Theta}_1]$.

**Hint**: [`np.diag`](https://numpy.org/doc/stable/reference/generated/numpy.diag.html)

```python
var_thetahat = ... # TODO
```

```python
grader.check("q3p2")
```

## 3.3 Compute the radiuses of the confidence intervals

We know that all of the parameters are Gaussian with variance given by the $p$'th diagonal entry in $\hat\sigma^2 \Sigma^{-1}_X$. We denote this variance as $v_p$. Then $v_p$ is the $p$-th entry in `var_thetahat`. 

Compute the radiuses of the 95\% confidence intervals for the slope parameters $\underline{\hat\theta}_1$ with:

$$\rho =\sqrt{Var[\underline{\hat\Theta}_1] } \left| F^{-1}_{\mathcal{N}}\left( \frac{1-\gamma}{2}\right)\right| $$



```python
gamma = 0.95
rho = ... # TODO
```

```python
grader.check("q3p3")
```

## 3.4 Tag as "significant" those parameters whose confidence interval does not include zero.

Create a 1D NumPy boolean array of the same size as `theta1hat` called `significant`. The $i$'th entry of `significant` should be `True` if the 95\% confidence interval for the corresponding slope parameter **does not** include 0, and `False` otherwise. In other words, an input is considered significant if its slope parameter is non-zero with a high degree of confidence. 

```python
significant = ...
```

```python
grader.check("q3p4")
```

## 3.5 Parameters table (done already)
Make a DataFrame with one row for each input. The index of the table should be the input names. The columns should be:
+ `slope`: the estimates of the slope parameter associated with the input. 
+ `slope stddev`: the standard deviation of the slope parameter.
+ `significant`: whether the input is significant according to part 3.4.

```python
params_table = pd.DataFrame(index=inputs,
             data={'slope':theta1hat,
                   'slope stddev':np.sqrt(var_thetahat),
                   'significant':significant})

params_table
```

## 3.6 Build an array of significant inputs

Extract the names of the significant inputs (i.e. inputs) from `params_table` using the `significant` array from part 3.4.

Store these significant input names as `significant_inputs`. `significant_inputs` should be a NumPy array with shape `(15,)`.

Here's one way you can do this that doesn't involve "for" loops:
1. Use the `significant` column to select the rows of the table corresponding to significant inputs. 
2. Use `.index` to obtain the names of the inputs for those rows. 
3. Use `.to_numpy()` to convert the result to a NumPy array.

```python
...
significant_inputs = ...  # TODO
```

```python
grader.check("q3p6")
```

## 3.7 Create a new table with significant inputs only (done already)

This table is called `data` and the target variable is now called `Y`.

```python
data = data3[significant_inputs].copy()
data['Y'] = data3['target_deathrate']
data
```

---

# Part 4: Feature subset selection 

## 4.1 Split `data` into training, validation, and testing datasets

We will use 70% of the data for training, 15% for validation, and 15% for testing.

1. Define `Dtrain` as the first `Ntrain` rows of `data`.
2. Define `Dvalidate` as the next `Nvalidate` rows of `data`.
3. Define `Dtest` as the last `Ntest` rows of `data`.

Here we use pandas' [iloc](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.iloc.html) method for selecting the three datasets.

```python
Ntrain = round(0.7*N)
Nvalidate = round(0.15*N)
Ntest = N - Ntrain - Nvalidate
Ntrain, Nvalidate, Ntest

Dtrain = data.iloc[:Ntrain,:]
Dvalidate = data.iloc[Ntrain:Ntrain+Nvalidate,:]
Dtest = data.iloc[Ntrain+Nvalidate:,:]
```

## 4.2 Linear regression training function

Create a function called `train` that receives a list of features `S` and a dataset `Dtrain` and does the following:
1. Selects the features `S` from `D` and stores them in `X`. (done already)
2. Selects the target values from `D` and stores them in `Y`. (done already)
3. Performs the linear regression calculations from parts 2.2, 2.3, and 2.4
4. Returns the estimated parameters $\hat\theta$

```python
def train(S, Dtrain):

    X = Dtrain[list(S)].values
    Y = Dtrain['Y'].values

    # 2.2 Center the inputs 
    muhatX = ...    # TODO
    Xc = ...    # TODO

    # 2.3 Center the outputs 
    muhatY = ...    # TODO
    Yc = ...    # TODO

    # 2.4 Compute the inverse input covariance matrix
    invCovX = ...    # TODO

    # 2.5 Compute the least squares estimates of the parameters
    theta1hat = ...    # TODO
    theta0hat = ...    # TODO

    return theta0hat, theta1hat
```

```python
theta0hat, theta1hat = train(['incidencerate','birthrate'], Dtrain)
```

```python
grader.check("q4p2")
```

## 4.3 Model evaluation function

Create a function called `perf` that receives the linear regression parameters `theta0hat` and `theta1hat`, their corresponding feature names `S`, and a dataset `D`, which may be the validation or the training dataset.

The function should evluate the means squared error (MSE) for the model using this data.

The steps are:
1. Select the features `S` from `D` and stores them in `X`. (done already)
2. Select the target values from `D` and stores them in `Y`. (done already)
3. Compute `Yhat`, as in part 2.6.
4. Use the MSE as the performance metric:

$$MSE = \frac{1}{N}\sum_{i=1}^{N} ( y_i-\hat{y}_i)^2$$

```python
def perf( S, theta0hat, theta1hat, D):

    X = D[list(S)].values
    Y = D['Y'].values

    # 2.6 Compute predictions for each of the samples
    Yhat = ... # TODO

    # 2.7 Model performance
    MSE = ... # TODO
    
    return MSE
```

```python
theta0hat, theta1hat = train(['incidencerate','birthrate'], Dtrain)
perf(['incidencerate','birthrate'], theta0hat, theta1hat, Dtrain)
```

```python
grader.check("q4p3")
```

## 4.4 Forward stepwise selection

Following is the code for forward stepwise feature selection that was described in the lab. When you run it, it will store the solution in variables with names `f_*`.

This part has no deliverables.

```python
P = len(significant_inputs)
curlyP = set(significant_inputs)
```

```python
curlyS = [set() for i in range(P+1)]
ellk = np.full(P+1,np.inf)

for k in range(1,P+1):
    
    curlyA = [set() for i in range(P-k+1)]
    ellkappa = np.full(P-k+1,np.inf)

    for kappa, phip in enumerate(curlyP-curlyS[k-1]):
        curlyA[kappa] = curlyS[k-1].union({phip})
        theta0hat, theta1hat = train( curlyA[kappa] , Dtrain)
        ellkappa[kappa] = perf(curlyA[kappa], theta0hat, theta1hat, 
                               Dvalidate)

    kappastar = ellkappa.argmin()
    curlyS[k] = curlyA[kappastar]
    ellk[k] = ellkappa[kappastar]

kstar = ellk.argmin()
Sstar = curlyS[kstar]
theta0star, theta1star = train(Sstar, Dtrain)
ellstar = perf(Sstar, theta0star, theta1star, Dtest)

# Store the results
f_ellk = ellk
f_ellstar = ellstar
f_kstar = kstar
```

## 4.5 Backward stepwise selection

Complete the code for backward stepwise selection.

```python
curlyS = [set() for i in range(P+1)]
curlyS[P] = set(significant_inputs)
ellk = np.full(P+1,np.inf)

for k in range(...):    # TODO

    curlyA = [set() for i in range(k+1)]
    ellkappa = np.full(k+1,np.inf)

    for kappa, phip in enumerate(...):    # TODO
        curlyA[kappa] = ...    # TODO
        theta0hat, theta1hat = train( curlyA[kappa] , Dtrain)
        ellkappa[kappa] = perf(curlyA[kappa], theta0hat, theta1hat, Dvalidate)

    kappastar = ellkappa.argmin()
    curlyS[k] = curlyA[kappastar]
    ellk[k] = ellkappa[kappastar]

kstar = ellk.argmin()
Sstar = curlyS[kstar]
theta0star, theta1star = train(Sstar , Dtrain)
ellstar = perf(Sstar, theta0star, theta1star, Dtest)

# Store the results
b_ellk = ellk
b_ellstar = ellstar
b_kstar = kstar
```

```python
grader.check("q4p5")
```

## Plot

The following plot shows the results of forward and backward stepwise feature selection for this regression problem. In each case, the star indicates the test error for the model with the smallest validation error.

```python
plt.figure(figsize=(10,5))

plt.plot(range(P+1),f_ellk,'o-',color='blue',label='forward',linewidth=3)
plt.plot([f_kstar,f_kstar],[f_ellk[f_kstar],f_ellstar],color='blue',linestyle='--',linewidth=2)
plt.plot(f_kstar,f_ellstar,'*',color='blue',markersize=22)

c = 'darkorange'
plt.plot(range(P+1),b_ellk,'o-',color=c,label='backward',linewidth=2)
plt.plot([b_kstar,b_kstar],[b_ellk[b_kstar],b_ellstar],color=c,linestyle=':',linewidth=2)
plt.plot(b_kstar,b_ellstar,'*',color=c,markersize=16)
plt.legend(fontsize=16)

plt.grid(linestyle=':')

# plt.ylim(340,480)
# plt.xlim(0,16)
plt.xticks(range(16),fontsize=16)
plt.xlabel('k',fontsize=16)
plt.ylabel('MSE',fontsize=16)
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1p2 (points: None)
```
>>> non_numerical_inputs == {'geography', 'binnedinc'}
True
```
```
>>> null_entry_inputs == {'pctsomecol18_24', 'pctprivatecoveragealone', 'pctemployed16_over'}
True
```


Tests for q1p3 (points: None)
```
>>> data1.shape[0] == 3047
True
```
```
>>> data1.shape[1] == 26
True
```
```
>>> A = data1.columns.to_numpy().copy()
>>> A.sort()
>>> h = [get_hash(a) for a in A]
>>> h == ['0ad6752935593b398c5b2e03e67956a8', '01915914efcc68fd29e1958648d4ad9a', '6e2ef10ca497cec351147c5548df65b0', 'bd6282c738ee849dcf6df51cbacb6fb6', '9aed79c74c3b56d88a469bd924d5536a', 'f6f19e7962a9cb1519a04973bb9e8767', '019ecaf6b6730586dfbed865d892a780', '426951735326c71c501d00ef63df3e71', '311b38f064bf47006677ad0580e442f8', 'ea892d7d8efd039319033bd345077f3b', '98af7560c20e755217434ab3ef503d00', 'b0b242280111a535759eeaece794e5cf', 'c7f0a021cd78750c0d124f83b0fdb731', '979395d25696c193bd3e11f9b7f625db', '5aa96cfedb0cd72cf87612147fab5aba', '3c2052cb47e7261583fabe6f76b74a64', '1d4775730357445b9b289574d854c3c5', 'f636665ff9d00921e2e0178931e00cf3', '46926b87727f040d06ae16b5f01f858e', 'c096e7dd49576ce7b9be6461f5c3650f', '0941e7200c731e193120b119d5e3be71', '7c85111de502c96aba2feafba12005d4', 'b533d326331997223720398741cb3f93', '9862bc7c6c3f3bf6b5bdd11faf2b18c2', 'eb4ff559f5f2be0b9d8ee50d9532703e', '102fb5f603a205bd350961efb6054b1c']
True
```


Tests for q1p4 (points: None)
```
>>> get_hash(thresh, 4) == '4fd32b31ea0748490ef09f7fe56b8d6f'
True
```
```
>>> get_hash(data2.shape[1], 2) == 'e237633a8b413a04dd5028207f03c1c2'
True
```


Tests for q1p5 (points: None)
```
>>> data3.shape[0] < 2900
True
```
```
>>> get_hash(data3.shape[0], 4) == 'c455fdadcc5232e7e119979863ef3848'
True
```


Tests for q1p6 (points: None)
```
>>> A = top_5_sort.copy()
>>> A.sort()
>>> h = [get_hash(a) for a in A]
>>> h == ['6e2ef10ca497cec351147c5548df65b0', '9aed79c74c3b56d88a469bd924d5536a', '426951735326c71c501d00ef63df3e71', 'c096e7dd49576ce7b9be6461f5c3650f', 'eb4ff559f5f2be0b9d8ee50d9532703e']
True
```


Tests for q2p1 (points: None)
```
>>> N > 1000 and N < 3000 and (D > 10) and (D < 30) and (inputs.shape[0] == 23) and (Y.shape[0] == 2895) and (X.shape == (2895, 23))
True
```
```
>>> get_hash(N, 4) == 'c455fdadcc5232e7e119979863ef3848'
True
```
```
>>> get_hash(D, 2) == '6bc071ec71e51c704acd13cdc898fd93'
True
```


Tests for q2p2 (points: None)
```
>>> get_hash(muhatX.shape[0], 2) == '6bc071ec71e51c704acd13cdc898fd93'
True
```
```
>>> get_hash(muhatX, 4) == 'd3b9e3451b30c05adf3eae1beb231ce9'
True
```
```
>>> get_hash(Xc.mean(axis=0), 4) == 'bb7104c2c3244ed0dcd24fd8b051d65e'
True
```


Tests for q2p3 (points: None)
```
>>> get_hash(muhatY, 4) == '7262c932665bb83af233f3ac9c4ab604'
True
```
```
>>> get_hash(Yc, 4) == '343cfcc5198ad67c6fcb742a3bbaad05'
True
```


Tests for q2p4 (points: None)
```
>>> invCovX.shape == (23, 23)
True
```
```
>>> get_hash(invCovX, 4) == '47fa6d90695b03c5d2ed8e12a701cdda'
True
```


Tests for q2p5 (points: None)
```
>>> theta1hat.shape[0] == 23 and theta0hat.shape == ()
True
```
```
>>> get_hash(theta1hat, 4) == '598d817c35ff8dd43907a6470ed5ebfa'
True
```
```
>>> get_hash(theta0hat, 4) == 'aa75626390499106698bbc42e492af3a'
True
```


Tests for q2p6 (points: None)
```
>>> get_hash(Yhat, 4) == '1f8b043cdf729e97b6f8b8592bed2336'
True
```


Tests for q2p7 (points: None)
```
>>> get_hash(R2, 4) == 'f71f388290a0e65c50c79a675b8f1811'
True
```


Tests for q3p1 (points: None)
```
>>> get_hash(sigmahat2, 4) == 'd4dbd5f2e721125a067d995f84a394f0'
True
```


Tests for q3p2 (points: None)
```
>>> var_thetahat.shape[0] == 23
True
```
```
>>> get_hash(var_thetahat, 4) == '53ec6c6ab5e14697ed4eed1de7877090'
True
```


Tests for q3p3 (points: None)
```
>>> get_hash(rho, 4) == '68e52ca65ae047457e67cb4d8d3b53d8'
True
```


Tests for q3p4 (points: None)
```
>>> bool(np.sum(significant) == 15)
True
```
```
>>> get_hash(rho, 4) == '68e52ca65ae047457e67cb4d8d3b53d8'
True
```


Tests for q3p6 (points: None)
```
>>> type(significant_inputs) == np.ndarray
True
```
```
>>> significant_inputs.shape == (15,)
True
```
```
>>> A = significant_inputs.copy()
>>> A.sort()
>>> [get_hash(a) for a in A] == ['01915914efcc68fd29e1958648d4ad9a', '6e2ef10ca497cec351147c5548df65b0', '426951735326c71c501d00ef63df3e71', 'ea892d7d8efd039319033bd345077f3b', '98af7560c20e755217434ab3ef503d00', 'b0b242280111a535759eeaece794e5cf', 'c7f0a021cd78750c0d124f83b0fdb731', '979395d25696c193bd3e11f9b7f625db', '5aa96cfedb0cd72cf87612147fab5aba', '3c2052cb47e7261583fabe6f76b74a64', '1d4775730357445b9b289574d854c3c5', '46926b87727f040d06ae16b5f01f858e', 'c096e7dd49576ce7b9be6461f5c3650f', 'b533d326331997223720398741cb3f93', '9862bc7c6c3f3bf6b5bdd11faf2b18c2']
True
```


Tests for q4p2 (points: None)
```
>>> (theta0hat, theta1hat) = train(['incidencerate', 'birthrate'], Dtrain)
>>> get_hash(theta0hat, 4) == 'e2965e9ecce803a6cb50701704ebe38d'
True
```
```
>>> (theta0hat, theta1hat) = train(['incidencerate', 'birthrate'], Dtrain)
>>> get_hash(theta1hat, 4) == '1d6832240b73566e03a4145fd47d03fb'
True
```


Tests for q4p3 (points: None)
```
>>> (theta0hat, theta1hat) = train(['incidencerate', 'birthrate'], Dtrain)
>>> get_hash(theta1hat, 4) == '1d6832240b73566e03a4145fd47d03fb' and get_hash(theta0hat, 4) == 'e2965e9ecce803a6cb50701704ebe38d'
True
```


Tests for q4p5 (points: None)
```
>>> get_hash(b_ellstar, 4) == '2a328cccc9fe2d321a6da1cd250fabf2'
True
```
```
>>> get_hash(b_kstar, 2) == '02ad845d2f8ac57a5c265ea083b2169b'
True
```
