```python
# Initialize Otter
import otter
grader = otter.Notebook("jupyter_demo.ipynb")
```

---

<h3><center>Python Bootcamp for Engineers</center></h3>

<h2><center>Demo for Jupyter Notebooks based exercises <br></center></h2>

---

```python
from resources.hashutils import *
```

# Part 1: A simple calculation
Compute the area of a square with sides of length L=2 meters

```python
L = 2
area = ...
```

```python
grader.check("p1")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for p1 (points: None)
```
>>> area > 0 and area < 10
True
```
```
>>> get_hash(area, 3) == '07078a97d66756f213dbca3e379bf084'
True
```
