```python
# Initialize Otter
import otter
grader = otter.Notebook("lab4.ipynb")
```

---

<h1><center>SDSE Lab 4 <br><br> Linear regression and Feature selection </center></h1>

---

In this lab we will use linear regression to predict cancer mortality rates based on data obtained from the American Community Survey of the [U.S. Census Bureau](https://www.census.gov/). The lab has five parts. In part 1 you will load the data and do basic manipulations using [pandas](https://pandas.pydata.org/docs/index.html). Pandas is a Python package that specializes in tabular data. It is widely used in data science and machine learning since the data in these fields are usually structured as a table. Pandas is a very powerful library that is well worth investing some time in. [Here](https://pandas.pydata.org/docs/getting_started/index.html#getting-started) and [here](https://pandas.pydata.org/docs/user_guide/index.html) are resources to learn more.

In part 2 you will perform linear regression on the full feature set. In part 3 you will assess the performance of the linear regression model using the coefficient of determination. Then, in part 4 you will compute confidence intervals for the slope parameters of the linear regression model. Finally, in part 5 you will run the forward and backward stepwise feature selection algorithms and estimate the performance of the resulting model using a test dataset.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
import scipy.stats as stats
from resources.hashutils import *
```

# Part 1:  Load and clean the data

## 1.1 Load the data into a pandas DataFrame

Use [pd.read_csv](https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html) to load the data from `cancerdata.csv`.

You can obtain information about the data using these DataFrame methods and attributes:
+ [`data0.head()`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.head.html): displays the first 5 rows of the DataFrame.
+ [`data0.tail()`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.tail.html): displays the last 5 rows of the DataFrame.
+ [`data0.shape`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.shape.html): a tuple with the number of rows and of columns in the DataFrame.
+ [`data0.columns`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.columns.html): the column headers.
+ [`data0.index`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.index.html): a unique identifier for each row.

```python
data0 = ...  # TODO
```

## 1.2 Inspect the columns

Run `data0.info()` and note:
 a) which inputs are non-numerical (Dtype=object), and
 b) which inputs have null entries (Non-Null Count<3047).

Store the names ('Column' entry) of the non-numerical inputs in a [set](https://www.w3schools.com/python/python_sets.asp) called `non_numerical_inputs`. Store the names of inputs with null entries in a set called `null_entry_inputs`.

**Note**: If not all of the rows of `data0.info()` are displayed, you'll probably have this message at the bottom:

*``Output is truncated. View as a scrollable element or open in a text editor. Adjust cell output settings...''*

```python
data0.info()
```

```python
non_numerical_inputs = {...,...}      # TODO
null_entry_inputs = {...,...,...}     # TODO
```

```python
grader.check("q1p2")
```

## 1.3 Discard non-numerical columns

Remove the two columns with non-numeric data.

Hints:
+ `data0.dtypes` lists the data types for each column.
+ You can construct a boolean mask for non-numeric columns with `data0.dtypes=='object'`.
+ Use that mask to index `data0.columns`
+ Use [`data0.drop`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.drop.html) to remove the selected columns.
+ Save the result as `data1`

```python
ind = ...               # TODO
drop_cols = ...         # TODO
data1 = data0.drop(...)       # TODO
```

```python
grader.check("q1p3")
```

## 1.4 Discard columns where more than 10% of values are nan

Hints:
+ [`.dropna`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.dropna.html)
+ Use `axis=1` to drop columns (as opposed to `axis=0` for rows.
+ Use the `thresh` argument. The condition for dropping a column is that it has less than `round(0.9*data1.shape[0])` non-nans.
+ Save the result as `data2`

```python
thresh = ...                # TODO
data2 = data1.dropna(...)     # TODO
```

```python
grader.check("q1p4")
```

## 1.5 Drop all rows that contain one or more nans.

Use the [`dropna`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.dropna.html) method of the `data2` DataFrame to remove rows with one or more nans.

Save the result as `data3`.

```python
data3 = ...      # TODO
```

```python
grader.check("q1p5")
```

## 1.6 Inspect correlations

Next we'll look at the sample correlation coefficients between each of the inputs and the output (a.k.a. target variable) `target_deathrate`. This is a quick way to check which of the inputs may be most useful to include in a model. Correlations only provide an initial guess, however. Remember that the correlation coefficient only measures the linear relationship between variables. That's perfect when the model is linear (as in this lab activity), but less useful for nonlinear models.

1) Use `data3.corr()` to build the correlations matrix.
2) Inspect the column (or row) corresponding to `target_deathrate`.
3) Rank (i.e. sort) the inputs from most to least correlated with the output. This ranking is in terms of the absolute value of the sample correlation coefficient.
4) Save the top 5 correlated inputs to `top_5_sort`. `top_5_sort` should be a numpy array with shape `(5,)`.

Hints:
+ [`abs`](https://pandas.pydata.org/docs/reference/api/pandas.Series.abs.html)
+ [`sort_values`](https://pandas.pydata.org/docs/reference/api/pandas.Series.sort_values.html)
+ [`to_numpy`](https://pandas.pydata.org/docs/reference/api/pandas.Series.to_numpy.html)

```python
# correlation matrix
C = ...         # TODO

# vector correlations between the inputs and target_deathrate
corr_target = ...         # TODO

# sorted corr_target_sort
corr_target_sort = ...         # TODO

# top 5 correlations with target_deathrate
top_5_sort = ...         # TODO
```

```python
grader.check("q1p6")
```

## 1.7 Make a scatter plot

Make a scatter plot of the data with the most correlated input along the x axis, and the target along the y axis.

Hint: You can use the plotting function attached to the DataFrame: [data3.plot(kind='scatter',x=..., y=...)](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.plot.html)

```python
data3.plot(kind='scatter',x=...,y=...)   # TODO
```

---

# Part 2: Solve linear regression

## 2.1 Extract `X` and `Y` from `data3`

The next cell extracts the `X` and `Y` matrices from `data3` and constructs a list of `inputs`. Define the number of samples `N` and the total number of inputs `D`.

```python
X = data3.drop(columns='target_deathrate').values
all_inputs = data3.columns.values
all_inputs = all_inputs[all_inputs!='target_deathrate']
Y = data3['target_deathrate'].values

N = ...
D = ...
```

```python
grader.check("q2p1")
```

## 2.2 Construct $\mathbb{X}$ defined in Eq. (6.112) of the reader

Implement the formula and store the result in the variable name `Xe`.

**Hint**: 
+ [np.hstack](https://numpy.org/doc/2.1/reference/generated/numpy.hstack.html)
+ [np.ones](https://numpy.org/doc/2.3/reference/generated/numpy.ones.html)

```python
Xe = ...
```

```python
grader.check("q2p2")
```

## 2.3 Find the solution $\underline{\hat\theta}$ of linear regression using Eq. (6.121) of the reader.

Implement the formula and store the result in the variable name `thetahat`.

Hint: 
+ [np.linalg.inv](https://numpy.org/doc/2.3/reference/generated/numpy.linalg.inv.html)
+ Matrix multiplication in NumPy can be achieved with the `@` operator or with [np.matmul](https://numpy.org/doc/2.3/reference/generated/numpy.matmul.html)

```python
thetahat = ...
```

```python
grader.check("q2p3")
```

## 2.4 Extract $\hat\theta_0$ and $\underline{\hat\theta}_1$ from $\underline{\hat\theta}$ as per Eq. (6.113) of the reader.

Eq. (6.113) shows the arrangement of candidate parameters $\theta_0$ and $\underline{\theta}_1$ into a candidate parameter vector $\underline{\theta}$. The same arrangement applies to the parameter estimates $\hat\theta_0$ and $\underline{\hat\theta}_1$. Unpack $\hat\theta$ into $\hat\theta_0$ and $\underline{\hat\theta}_1$, and save these respectively to variables `theta0hat` and `theta1hat`.


```python
theta0hat = ...
theta1hat = ...
```

```python
grader.check("q2p4")
```

---

# Part 3: Evaluate model performance

## 3.1 Compute predictions for each of the training samples using Eq. (6.111)

Implement the formula and store the result in the variable name `Yhat`.

```python
Yhat = ... 
```

```python
grader.check("q3p1")
```

## 3.2 Compute the coefficient of determination using Eq. (5.21). 

Implement the formula and store the result in the variable name `R2`.

```python
R2 = ...   
```

```python
grader.check("q3p2")
```

---

# Part 4: Quantify parameter uncertainty

## 4.1 Compute the average input value using Eq. (6.117) of the reader

Implement the formula and store the result in the variable name `muhatX`.

Hint: You can use the `axis` argument of [np.mean](https://numpy.org/doc/2.2/reference/generated/numpy.mean.html) to take an average across all samples. 

```python
muhatX = ...
```

```python
grader.check("q4p1")
```

## 4.2 Compute the centered inputs using Eq. (6.129) 


Hint: 
+ The formula for `Xc` has $\mathbf{1}_N\hat\mu_X$. The broadcasting rules of numpy make multiplying $\hat\mu_X$ by $\mathbf{1}_N$ unnecessary.
+ Check that the column-wise means of `Xc` equal zero (to machine precision). 

```python
Xc = ...
```

```python
grader.check("q4p2")
```

## 4.3 Compute $\hat\sigma^2$, an unbiased estimate $\sigma^2$

Equation (6.105) of the reader gives the formula for estimating the variance of the measurement noise $\mathcal{E}$ when there is only one input ($D=1$). 

$$\hat\sigma^2 = \frac{1}{N-2} \sum_{i=1}^{N}(y_i-\hat{y}_i)^2$$

The generalization of this formula to the case of $D$ inputs is,

$$\hat\sigma^2 = \frac{1}{N-D-1} \sum_{i=1}^{N}(y_i-\hat{y}_i)^2$$

Implement this formula and assign the result to the variable `sigmahat2`.

```python
sigmahat2 = ...
```

```python
grader.check("q4p3")
```

## 4.4 Compute the variances of the slope parameters

The slope parameters are contained in the array $\underline{\hat\theta}_1$. Compute their covariance matrix using Eq. (6.165) of the reader. In this equation you should replace the true noise level $\sigma^2$ with the estimate $\hat\sigma^2$ calculated in part 4.3. Save this matrix to the variable `CovThetaHat1`. Then extract the diagonal entries of this matrix, which correspond to the variances of each of the slope parameters. Save this to a 1D array called `varThetaHat1` 

**Hint**: [`np.diag`](https://numpy.org/doc/stable/reference/generated/numpy.diag.html)

```python
CovThetaHat1 = ...
varThetaHat1 = ...
```

```python
grader.check("q4p4")
```

## 4.5 Compute the radiuses of confidence intervals on the slope parameters

Having the variances of each of the slope parameters in `varThetaHat1`, and assuming that the output measurement noise is Gaussian, we can apply our generic formula from lecture for the radius of a confidence interval. This gives the following for the radius of a confidence interval for $\theta_d$:
$$\rho_d =\sqrt{v_d} \left| F^{-1}_{\mathcal{N}}\left( \frac{1-\gamma}{2}\right)\right| $$
Here $v_d$ is the variance of the $d$'th slope parameter, i.e. `varThetaHat1[d]`.

Compute the radiuses of 95\% confidence intervals for each of the slope parameters. Store these in the array `rho`.

```python
gamma = 0.95

rho = ... 
```

```python
grader.check("q4p5")
```

## 4.6 Tag as "significant" those parameters whose confidence interval does not include zero.

Create a 1D NumPy boolean array of the same size as `theta1hat` called `significant`. The $d$'th entry of `significant` should be `True` if the 95\% confidence interval for the corresponding slope parameter **does not** include 0, and `False` otherwise. In other words, an input is considered significant if its slope parameter is non-zero with 95\% confidence. 

```python
significant = ...
```

```python
grader.check("q4p6")
```

## 4.7 Parameters table (done already)
Make a DataFrame with one row for each input. The index of the table should be the input names. The columns should be:
+ `slope`: the estimates of the slope parameter associated with the input. 
+ `slope stddev`: the standard deviation of the slope parameter.
+ `significant`: whether the input is significant according to part 3.4.

```python
params_table = pd.DataFrame(index=all_inputs,
             data={'slope':theta1hat,
                   'slope stddev':np.sqrt(varThetaHat1),
                   'significant':significant})

params_table
```

## 4.8 Build an array of significant inputs

Extract the names of the significant inputs from `params_table` using the `significant` array from part 4.7.

Store these significant input names as `significant_inputs`. 

`significant_inputs` should be a NumPy array with shape `(15,)`.

Here's one way you can do this that doesn't require a "for" loop:
1. Use the `significant` column to select the rows of the table corresponding to significant inputs. 
2. Use `.index` to obtain the names of the inputs for those rows. 
3. Use `.to_numpy()` to convert the result to a NumPy array.

```python
...
significant_inputs = ...
```

```python
grader.check("q4p8")
```

## 4.9 Create a new table with significant inputs only (done already)

This table is called `data` and the target variable is now called `Y`.

```python
data = data3[significant_inputs].copy()
data['Y'] = data3['target_deathrate']
data
```

---

# Part 5: Feature selection 

## 5.1 Split `data` into training, validation, and testing datasets (done already)

We will use 70% of the data for training, 15% for validation, and 15% for testing.

1. Define `Dtrain` as the first `Ntrain` rows of `data`.
2. Define `Dvalidate` as the next `Nvalidate` rows of `data`.
3. Define `Dtest` as the last `Ntest` rows of `data`.

Here we use pandas' [iloc](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.iloc.html) method for selecting the three datasets.

```python
Ntrain = round(0.7*N)
Nvalidate = round(0.15*N)
Ntest = N - Ntrain - Nvalidate
Ntrain, Nvalidate, Ntest

Dtrain = data.iloc[:Ntrain,:]
Dvalidate = data.iloc[Ntrain:Ntrain+Nvalidate,:]
Dtest = data.iloc[Ntrain+Nvalidate:,:]
```

## 5.2 Linear regression training function

Create a function called `train` that receives a list of features `S` and a dataset `Dtrain` and does the following:
1. Selects the features `S` from `Dtrain` and stores them in `X`. (done already)
2. Selects the target values from `Dtrain` and stores them in `Y`. (done already)
3. Performs the linear regression calculations from parts 2.2 and 2.3 (copy your code from those parts into the `train` method)
4. Returns the estimated parameter vector `thetahat` containing $\hat\theta_0$ and $\underline{\hat\theta}_1$

```python
def train(S, Dtrain):

    X = Dtrain[list(S)].values
    Y = Dtrain['Y'].values
    N = X.shape[0]

    # 2.2 Construct $\mathbb{X}$
    Xe = ...

    # 2.3 Compute thetahat
    thetahat = ...

    return thetahat
```

```python
# Use this cell to test your code

thetahat = train(['incidencerate','birthrate'], Dtrain)
```

```python
grader.check("q5p2")
```

## 5.3 Model evaluation function

Create a function called `eval` that receives
+ `S`... a set of input feature names

 the linear regression parameters `thetahat`, the corresponding feature names `S`, and a dataset `D`, which may be the validation, the training, or the testing dataset.

The function should evaluate the mean squared error (MSE) of the model using this data.

The steps are:
1. Select the features `S` from `D` and stores them in `X`. (done already)
2. Select the target values from `D` and stores them in `Y`. (done already)
3. Compute `Yhat`, as in part 3.1.
4. Evaluate $R^2$.

```python
def train_eval(S, Dtrain, Deval):

    # fix the order of the inputs S by casting it to a list
    S = list(S)

    # train linear regression (call the training function you already made)
    thetahat = ...

    # unpack the linear regression coefficients
    theta0hat = ...
    theta1hat = ...

    # unpack the evaluation data
    X = Deval[S].values
    Y = Deval['Y'].values

    # compute predictions for each of the samples (same as part 3.1)
    Yhat = ...

    # 2.7 Evaluate performance with R2 (same as part 3.2)
    R2 = ...
    
    return R2
```

```python
# Use this cell to test your code
r2 = train_eval(['incidencerate','birthrate'], Dtrain, Dvalidate)
```

```python
grader.check("q5p3")
```

## 5.4 Forward feature selection (done already)

Below is a method that implements the forward feature selection algorithm that was described by your GSI in lab. 

This part has no deliverables.

```python
def forward_selection(all_inputs):

    # allocate stuff
    P = len(all_inputs)                     # ... number of inputs
    setF = set(all_inputs)                  # ... fixed set of all inputs
    setFk = [set() for i in range(P+1)]     # ... array of best input set at each stage
    ellk = np.empty(P+1)             # ... array of best performance at each stage

    # initialize
    setFk[0] = set()
    ellk[0] = train_eval(set(),Dtrain,Dvalidate)

    # loop through stages
    for k in range(1,P+1):
        
        # allocate for inner loop 
        setA = [set() for i in range(P-k+1)]  # ... array of input sets to evaluate
        pkappa = np.full(P-k+1,np.inf)        # ... array of performance values

        # inner loop: through all remaining inputs to add
        for kappa, phip in enumerate(setF-setFk[k-1]):

            # add the p-th input (number kappa among remaining inputs)
            setA[kappa] = setFk[k-1].union({phip})

            # train with the training data, evaluate with validation data
            pkappa[kappa] = train_eval(setA[kappa],Dtrain,Dvalidate)

        # keep the best set and its performance.
        kappastar = pkappa.argmax()
        setFk[k] = setA[kappastar]
        ellk[k] = pkappa[kappastar]

    # keep the best over all stages
    kstar = ellk.argmax()
    Fstar = setFk[kstar]

    # save the best parameters and their performance
    ellstar = train_eval(Fstar, Dtrain, Dtest)

    return ellk, ellstar, kstar
```

```python
# You can run forward selection on the collection of significant inputs

f_ellk, f_ellstar, f_kstar = forward_selection(significant_inputs)
```

## 5.5 Backward feature removal

Implement backward feature removal.

```python
def backward_removal(all_inputs):

    ...

    return ellk, ellstar, kstar
```

```python
# Use this cell to test your code
b_ellk, b_ellstar, b_kstar = backward_removal(significant_inputs)

b_ellk, b_ellstar, b_kstar
```

```python
grader.check("q5p5")
```

## 5.6 Run forward and backward selection (done already)

```python
f_ellk, f_ellstar, f_kstar = forward_selection(significant_inputs)
b_ellk, b_ellstar, b_kstar = backward_removal(significant_inputs)
P = len(significant_inputs)
```

## 5.7 Plot (done already)

```python
plt.figure(figsize=(10,5))

# Plot forward seleection in blue
color = 'blue'
plt.plot(range(P+1),f_ellk,'o-',color=color,linewidth=3,label='forward (validation)')
plt.plot([f_kstar,f_kstar],[f_ellk[f_kstar],f_ellstar],color=color,linestyle='--',linewidth=2)
plt.plot(f_kstar,f_ellstar,'*',color=color,markersize=22,label='forward (test)')

# Plot backward removal in orange
color = 'darkorange'
plt.plot(range(P+1),b_ellk,'o-',color=color,linewidth=2,label='backward (validation)')
plt.plot([b_kstar,b_kstar],[b_ellk[b_kstar],b_ellstar],color=color,linestyle=':',linewidth=2)
plt.plot(b_kstar,b_ellstar,'*',color=color,markersize=16,label='backward (test)')

plt.legend(fontsize=14)
plt.grid(linestyle=':')
plt.xticks(range(16),fontsize=16)
plt.xlabel('k (number of inputs included)',fontsize=16)
plt.ylabel(r'$R^2$',fontsize=16,rotation=0,labelpad=20)
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1p2 (points: None)
```
>>> non_numerical_inputs == {'geography', 'binnedinc'}
True
```
```
>>> null_entry_inputs == {'pctsomecol18_24', 'pctprivatecoveragealone', 'pctemployed16_over'}
True
```


Tests for q1p3 (points: None)
```
>>> data1.shape[0] > 3000 and data1.shape[0] < 3100
True
```
```
>>> data1.shape[1] == 26
True
```
```
>>> A = data1.columns.to_numpy().copy()
>>> A.sort()
>>> h = [get_hash(a) for a in A]
>>> h == ['0ad6752935593b398c5b2e03e67956a8', '01915914efcc68fd29e1958648d4ad9a', '6e2ef10ca497cec351147c5548df65b0', 'bd6282c738ee849dcf6df51cbacb6fb6', '9aed79c74c3b56d88a469bd924d5536a', 'f6f19e7962a9cb1519a04973bb9e8767', '019ecaf6b6730586dfbed865d892a780', '426951735326c71c501d00ef63df3e71', '311b38f064bf47006677ad0580e442f8', 'ea892d7d8efd039319033bd345077f3b', '98af7560c20e755217434ab3ef503d00', 'b0b242280111a535759eeaece794e5cf', 'c7f0a021cd78750c0d124f83b0fdb731', '979395d25696c193bd3e11f9b7f625db', '5aa96cfedb0cd72cf87612147fab5aba', '3c2052cb47e7261583fabe6f76b74a64', '1d4775730357445b9b289574d854c3c5', 'f636665ff9d00921e2e0178931e00cf3', '46926b87727f040d06ae16b5f01f858e', 'c096e7dd49576ce7b9be6461f5c3650f', '0941e7200c731e193120b119d5e3be71', '7c85111de502c96aba2feafba12005d4', 'b533d326331997223720398741cb3f93', '9862bc7c6c3f3bf6b5bdd11faf2b18c2', 'eb4ff559f5f2be0b9d8ee50d9532703e', '102fb5f603a205bd350961efb6054b1c']
True
```


Tests for q1p4 (points: None)
```
>>> thresh > 2700 and thresh < 2800
True
```
```
>>> data2.shape[1] == 24
True
```
```
>>> get_hash(thresh, 4) == '4fd32b31ea0748490ef09f7fe56b8d6f'
True
```
```
>>> A = data2.columns.to_list()
>>> A.sort()
>>> [get_hash(a) for a in A] == ['0ad6752935593b398c5b2e03e67956a8', '01915914efcc68fd29e1958648d4ad9a', '6e2ef10ca497cec351147c5548df65b0', 'bd6282c738ee849dcf6df51cbacb6fb6', '9aed79c74c3b56d88a469bd924d5536a', 'f6f19e7962a9cb1519a04973bb9e8767', '019ecaf6b6730586dfbed865d892a780', '426951735326c71c501d00ef63df3e71', '311b38f064bf47006677ad0580e442f8', 'ea892d7d8efd039319033bd345077f3b', '98af7560c20e755217434ab3ef503d00', 'b0b242280111a535759eeaece794e5cf', 'c7f0a021cd78750c0d124f83b0fdb731', '979395d25696c193bd3e11f9b7f625db', '5aa96cfedb0cd72cf87612147fab5aba', '3c2052cb47e7261583fabe6f76b74a64', '1d4775730357445b9b289574d854c3c5', '46926b87727f040d06ae16b5f01f858e', 'c096e7dd49576ce7b9be6461f5c3650f', '7c85111de502c96aba2feafba12005d4', 'b533d326331997223720398741cb3f93', '9862bc7c6c3f3bf6b5bdd11faf2b18c2', 'eb4ff559f5f2be0b9d8ee50d9532703e', '102fb5f603a205bd350961efb6054b1c']
True
```


Tests for q1p5 (points: None)
```
>>> data3.shape == (2895, 24)
True
```
```
>>> get_hash(data3.index, 4) == '5e33b407227c7dda3c52b4b51485be0e'
True
```


Tests for q1p6 (points: None)
```
>>> corr_target_sort.shape == (24,)
True
```
```
>>> A = top_5_sort.copy()
>>> A.sort()
>>> h = [get_hash(a) for a in A]
>>> h == ['6e2ef10ca497cec351147c5548df65b0', '9aed79c74c3b56d88a469bd924d5536a', '426951735326c71c501d00ef63df3e71', 'c096e7dd49576ce7b9be6461f5c3650f', 'eb4ff559f5f2be0b9d8ee50d9532703e']
True
```


Tests for q2p1 (points: None)
```
>>> N > 1000 and N < 3000 and (D > 10) and (D < 30) and (all_inputs.shape[0] == 23) and (Y.shape[0] == 2895) and (X.shape == (2895, 23))
True
```
```
>>> get_hash(N, 4) == 'c455fdadcc5232e7e119979863ef3848'
True
```
```
>>> get_hash(D, 2) == '6bc071ec71e51c704acd13cdc898fd93'
True
```


Tests for q2p2 (points: None)
```
>>> Xe.shape == (2895, 24)
True
```
```
>>> get_hash(Xe, 4) == '794bba9d89c4411be99153fc08634c78'
True
```


Tests for q2p3 (points: None)
```
>>> thetahat.shape == (24,)
True
```
```
>>> get_hash(thetahat, 4) == '20d4f19fa78ee48a1f7184205af8286b'
True
```


Tests for q2p4 (points: None)
```
>>> bool(theta0hat > 150 and theta0hat < 200)
True
```
```
>>> theta1hat.shape == (23,)
True
```
```
>>> bool(min(theta1hat) > -2 and max(theta1hat) < 2)
True
```
```
>>> get_hash(theta0hat, 4) == 'aa75626390499106698bbc42e492af3a'
True
```
```
>>> get_hash(theta1hat, 4) == '598d817c35ff8dd43907a6470ed5ebfa'
True
```


Tests for q3p1 (points: None)
```
>>> Yhat.shape == (2895,)
True
```
```
>>> get_hash(Yhat, 4) == '1f8b043cdf729e97b6f8b8592bed2336'
True
```


Tests for q3p2 (points: None)
```
>>> bool(R2 > 0.5 and R2 < 0.6)
True
```
```
>>> get_hash(R2, 4) == 'f71f388290a0e65c50c79a675b8f1811'
True
```


Tests for q4p1 (points: None)
```
>>> muhatX.shape == (23,)
True
```
```
>>> get_hash(muhatX, 4) == 'd3b9e3451b30c05adf3eae1beb231ce9'
True
```


Tests for q4p2 (points: None)
```
>>> Xc.shape == (2895, 23)
True
```
```
>>> get_hash(Xc, 3) == '02f312181e07841edaa62e20c1c346c7'
True
```


Tests for q4p3 (points: None)
```
>>> bool(sigmahat2 > 300 and sigmahat2 < 400)
True
```
```
>>> get_hash(sigmahat2, 4) == 'd4dbd5f2e721125a067d995f84a394f0'
True
```


Tests for q4p4 (points: None)
```
>>> CovThetaHat1.shape == (23, 23)
True
```
```
>>> varThetaHat1.shape[0] == 23
True
```
```
>>> get_hash(varThetaHat1, 4) == '53ec6c6ab5e14697ed4eed1de7877090'
True
```


Tests for q4p5 (points: None)
```
>>> rho.shape == (23,)
True
```
```
>>> get_hash(rho, 4) == '68e52ca65ae047457e67cb4d8d3b53d8'
True
```


Tests for q4p6 (points: None)
```
>>> bool(np.sum(significant) == 15)
True
```
```
>>> get_hash(rho, 4) == '68e52ca65ae047457e67cb4d8d3b53d8'
True
```


Tests for q4p8 (points: None)
```
>>> type(significant_inputs) == np.ndarray
True
```
```
>>> significant_inputs.shape == (15,)
True
```
```
>>> A = significant_inputs.copy()
>>> A.sort()
>>> [get_hash(a) for a in A] == ['01915914efcc68fd29e1958648d4ad9a', '6e2ef10ca497cec351147c5548df65b0', '426951735326c71c501d00ef63df3e71', 'ea892d7d8efd039319033bd345077f3b', '98af7560c20e755217434ab3ef503d00', 'b0b242280111a535759eeaece794e5cf', 'c7f0a021cd78750c0d124f83b0fdb731', '979395d25696c193bd3e11f9b7f625db', '5aa96cfedb0cd72cf87612147fab5aba', '3c2052cb47e7261583fabe6f76b74a64', '1d4775730357445b9b289574d854c3c5', '46926b87727f040d06ae16b5f01f858e', 'c096e7dd49576ce7b9be6461f5c3650f', 'b533d326331997223720398741cb3f93', '9862bc7c6c3f3bf6b5bdd11faf2b18c2']
True
```


Tests for q5p2 (points: None)
```
>>> thetahat = train(['incidencerate', 'birthrate'], Dtrain)
>>> bool(np.all(np.isclose(thetahat, [92.08, 0.196, -0.245], atol=0.01)))
True
```
```
>>> thetahat = train(['pctnohs18_24', 'pcths18_24', 'pcths25_over'], Dtrain)
>>> get_hash(thetahat, 4) == '1a6cfc8fc80de11991203ec4858901f6'
True
```


Tests for q5p3 (points: None)
```
>>> r2 = train_eval(['incidencerate', 'birthrate'], Dtrain, Dtrain)
>>> bool(np.isclose(r2, 0.158, atol=0.01))
True
```
```
>>> r2 = train_eval(['pctnohs18_24', 'pcths18_24', 'pcths25_over'], Dtrain, Dtrain)
>>> get_hash(r2, 4) == '231a369a26afc8f32767ddd9f7d18ee9'
True
```


Tests for q5p5 (points: None)
```
>>> b_ellk, b_ellstar, b_kstar = backward_removal(significant_inputs)
>>> b_ellk.shape == (16,) and bool(b_ellstar > 0.3 and b_ellstar < 0.5)
True
```
```
>>> b_ellk, b_ellstar, b_kstar = backward_removal(significant_inputs)
>>> get_hash(b_ellstar, 4) == '72029615b5043c216ba3f757db9636b1'
True
```
```
>>> b_ellk, b_ellstar, b_kstar = backward_removal(significant_inputs)
>>> get_hash(b_kstar, 2) == '02ad845d2f8ac57a5c265ea083b2169b'
True
```
