```python
# Initialize Otter
import otter
grader = otter.Notebook("lab6.ipynb")
```

---

<h1><center>SDSE Lab 6 <br><br> Ensemble methods and hyperparameter tuning </center></h1>

---




In this lab exercise we will apply several classification models to a problem in astrophysics. The problem is described [here](https://satyam5120.medium.com/predicting-a-pulsar-star-using-different-machine-learning-algorithms-d22ee8fc71b4) and [here](https://www.kaggle.com/datasets/colearninglounge/predicting-pulsar-starintermediate). It consists in labeling observations of space objects as either pulsars or not pulsars, based on the properties of an 'integrated profile' and on the DM-SNR curve. The dataset has 8 feature columns:
1. Mean of the integrated profile.
2. Standard deviation of the integrated profile.
3. Excess kurtosis of the integrated profile.
4. Skewness of the integrated profile.
5. Mean of the DM-SNR curve.
6. Standard deviation of the DM-SNR curve.
7. Excess kurtosis of the DM-SNR curve.
8. Skewness of the DM-SNR curve.

Our goal is to choose a classification model from the ones covered in this class. The procedure will follow these steps:
1. Load the data.
2. Remove null values.
3. Find number of samples per class.
4. Extract a test dataset.
5. Build models:
    + Logistic regression
    + Random forest
    + AdaBoost
    + Gradient Boosted Trees
6. Select and evaluate a final model.

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
rng_seed = 2434
```

```python
def get_param_order(param_grid):
    param0 = [key for key, value in param_grid.items() if len(value) == 2]
    param1 = param_grid.keys() - param0
    return [param0[0],param1.pop()]

def unpack_gridsearch(gs,param_grid,folds):

    result = gs.cv_results_
    
    cols = ['mean_test_accuracy','mean_test_balanced_accuracy']
        
    param_order = get_param_order(param_grid)

    f0 = 'param_{}'.format(param_order[0])
    f1 = 'param_{}'.format(param_order[1])
    f0vals = param_grid[param_order[0]]
    f1vals = param_grid[param_order[1]]
    
    X = dict.fromkeys(f0vals)

    for f0val in f0vals:
        X[f0val] = pd.DataFrame(index=f1vals,columns=cols,dtype=float)

    for f0val in f0vals:
        for f1val in f1vals:
            ind = (result[f0]==f0val) & (result[f1]==f1val)
            for col in cols:
                X[f0val].loc[f1val,col] = float(result[col][ind][0])

    return X, gs.best_params_, gs.best_estimator_, gs.best_score_

def plot_grid_result(X,param_grid,log=True):

    param_order = get_param_order(param_grid)
    cols = ['mean_test_accuracy', 'mean_test_balanced_accuracy']
    labels = ['accuracy', 'balanced_accuracy']
    colors = ['r','b']
    linestyles = ['--', '-']

    f0vals = param_grid[param_order[0]]

    plt.figure(figsize=(14,7))
    for f0ind, f0val in enumerate(f0vals):
        for c, col in enumerate(cols):
            if log:
                plt.semilogx(X[f0val][col],color=colors[c],linestyle=linestyles[f0ind],marker='o',label='{} {}'.format(f0val,labels[c]))
            else:
                plt.plot(X[f0val][col],color=colors[c],linestyle=linestyles[f0ind],marker='o',label='{} {}'.format(f0val,labels[c]))
    plt.legend(fontsize=14)
    plt.grid(':')
    plt.xticks(fontsize=14)
    plt.yticks(fontsize=14)
    plt.show()
```

# 1. Load the data

+ Load the data file `pulsar_data.csv` into a Pandas dataframe using [`pd.read_csv`](https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html)
+ Save the column headers corresponding to feature names (all except the `target_class`) to the variable `feature_names` ([Hint](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.columns.html))

```python
rawdata = ...
feature_names = ...
```

```python
grader.check("q1")
```

# 2. Remove null values

Remove any feature columns of `rawdata` with more than zero null values. To check for null values you can use the [`isnull`](https://pandas.pydata.org/docs/reference/api/pandas.isnull.html) method on each feature column. To remove a column, use the [`drop`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.drop.html) method with the `columns` argument.

**Hint**: `np.any`

```python
for feature in feature_names:
    if ...:
        rawdata = rawdata.drop(...)
```

```python
grader.check("q2")
```

# 3. Number of samples per class
Find the number of data points in each of the two classes. Use these variable names:
+ `N0` for the number of class 0 samples, 
+ `N1` for the number of class 0 samples. 

Notice that there is a significant imbalance of negative and positive samples in this dataset. 

```python
N0 = ...
N1 = ...
```

```python
grader.check("q3")
```

# 4. Extract the test dataset

Use scikit-learn's `train_test_split` method to split `rawdata` into trainging and testing parts. Keep 90% for training and 10% for testing. Remember to pass `random_state=rng_seed` to `train_test_split` so that the result is repeatable.

```python
from sklearn.model_selection import train_test_split
Xtrain, Xtest, ytrain, ytest = ...
```

```python
grader.check("q4")
```

# 5. Logistic regression 

## 5.1. Training

In this part we will use grid search to choose the values of the hyperparameters of the logistic regression pipeline. Begin by creating a pipeline with a `StandardScaler` followed by `LogisticRegression` classifier. Pass these parameters to the contructor of the logistic regression classifier:
+ `solver`: `liblinear`
+ `random_state`: `rng_seed`

Then fit the model using the training data.

```python
from sklearn.linear_model import LogisticRegression
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import StandardScaler

model = Pipeline([
    ('scaler' , ... ) ,
    ('model' , ... )
])
model.fit(...)
```

```python
grader.check("q5p1")
```

## 5.2. Hyperparameter tuning with grid search

Run grid search (see documentation here: [`GridSearchCV`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.GridSearchCV.html)) with 3-fold cross validation. The parameters to search over are:

+ the regularization function `penalty`. Candidate values are `l1` and `l2`
+ the regularization weight `C`. Candidate values are `np.logspace(-3, 1, 10)`.

**Note**: For `GridSearchCV` to work with the pipeline, it needs to set the parameters of the model using their string names. Within the pipeline, the parameters of the model are prefixed with `model__`. For example, in this case we have 

```python
param_grid = {
    'model__penalty' : ['l1','l2'],
    'model__C' : np.logspace(-3, 1, 10) }
```

Pass the following parameters to the `GridSearchCV` contructor.
+ `param_grid=param_grid`
+ `scoring=['accuracy','balanced_accuracy']`,
+ `cv=cvfolds`,
+ `refit='balanced_accuracy'`

```python
from sklearn.model_selection import GridSearchCV
cvfolds = 3
param_grid = ...
gs = GridSearchCV(...)
gs = gs.fit(...)
```

```python
result_logreg, best_params_logreg, best_model_logreg, best_score_logreg = unpack_gridsearch(gs,param_grid,cvfolds)
```


`lab6_utils.py` contains useful functions for unpacking and plotting the result of the grid search. Use the `lab6.unpack_gridsearch` method to extract information from the grid search solution. This method returns:
1. A pandas dataframe with the cross-validated performance metrics for each point on the grid.
2. A dictionary with the best-case hyperparameter values.
3. The classifier corresponding to the best hyperparameter values.


`lab6.plot_grid_result` plots the results of the grid search.

```python
plot_grid_result(result_logreg,param_grid)
```

```python
grader.check("q5p2")
```

## 5.3. Package it into a function

Next we will repeat this process with several other classification models. To keep the code clean, we will first collect the steps into a single function. This function receives as inputs:
1. The classifier object, e.g. `LogisticRegression(solver='liblinear',random_state=rng_seed)`
2. The `param_grid` dictionary that defines the search space for `GridSearchCV`.

It should:
1. Create the pipeline model.
2. Construct the `GridSearchCV` object as was done in part 5.2
3. Run `fit` on the grid search object, using the training data (no need to pass the data as an argument to `build_grid_plot`; it's a global variable)
4. Run `unpack_gridsearch` to obtain `result`, `best_params`, `best_model`, and `best_score`.
5. Plot the result with `lab6.plot_grid_result`
6. `return result, best_params, best_model, best_score`


**Note**: This part has no autograder test. Any errors in the `build_grid_plot` function should be caught by tests in subsequent parts. 

```python
def build_grid_plot(model,param_grid):
    
    # 1. Create the pipeline model
    model = Pipeline([
    ...
    ...
    ])

    # 2. Construct the `GridSearchCV` object as was done in part 5.2
    gs = GridSearchCV(...)
    
    # 3. Run `fit` on the grid search object, using the training data (no need to pass the data as an argument to `build_grid_plot`; it's a global variable)
    gs = gs.fit(...)
    
    # 4. Run `unpack_gridsearch` to obtain `result`, `best_params`, `best_model`, and `best_score`.
    result, best_params, best_model, best_score = lab6.unpack_gridsearch(...)

    # 5. Plot the result with `lab6.plot_grid_result`
    lab6.plot_grid_result(...)

    # 6. `return result, best_params, best_model, best_score`
    return ...
```

# 6. Random forest

Use the following parameters for the hyper-parameter search:
+ `max_features: ['sqrt','log2']`
+ `n_estimators: np.linspace(2, 100, 5, dtype=int)`

**Note**: Remember to set the random state for the model in the model's constructor.

```python
from sklearn.ensemble import RandomForestClassifier

param_grid_rf = ...

result_rf, best_params_rf, best_model_rf, best_score_rf = build_grid_plot(...)
```

```python
grader.check("q6")
```

# 7. AdaBoost

Use the following parameters for the hyper-parameter search:
+ `learning_rate: [0.01,0.1]`
+ `n_estimators: np.linspace(20, 100, 5, dtype=int)`

**Note**: Remember to set the random state for the model in the model's constructor.

```python
from sklearn.ensemble import AdaBoostClassifier

param_grid_ab = ...

result_ab, best_params_ab, best_model_ab, best_score_ab = build_grid_plot(...)
```

```python
grader.check("q7")
```

# 8. Gradient Boosted Trees

Use the following parameters for the hyper-parameter search:
+ `learning_rate: [0.1,1.0]`
+ `n_estimators: [50,75,100]`

**Note**: Remember to set the random state for the model in the model's constructor.

```python
from sklearn.ensemble import GradientBoostingClassifier

param_grid_gbc = ...

result_gbc, best_params_gbc, best_model_gbc, best_score_gbc = build_grid_plot(...)
```

```python
grader.check("q8")
```

# 9. Final model selection

We have now build 4 separate classifiers: logistic regression, random forest, AdaBoost, and Gradient boosted trees. These are stored below in the `all_models` list. Select from this ist the classifier with the best score. Save the corresponding name, model, and score to variables `best_name`, `best_model`, and `best_score` respoenctively. 

Then compute the test accuracy and balanced accuracy for the selected model, using its `predict` function and the `accuracy_score` and `balanced_accuracy_score` functions. Store these as `test_accuracy` and `test_balanced_accuracy`.

```python
from sklearn.metrics import balanced_accuracy_score, accuracy_score

all_models = [
    ('logreg',best_model_logreg, best_score_logreg),
    ('rf',best_model_rf, best_score_rf),
    ('ab',best_model_ab, best_score_ab),
    ('gbc',best_model_gbc, best_score_gbc),
]

...
best_name = ...
best_model = ...
best_score = ...

ypred = ...
test_accuracy = ...
test_balanced_accuracy = ...
```

```python
grader.check("q9")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> set(feature_names) == {' Mean of the integrated profile', ' Standard deviation of the integrated profile', ' Excess kurtosis of the integrated profile', ' Skewness of the integrated profile', ' Mean of the DM-SNR curve', ' Standard deviation of the DM-SNR curve', ' Excess kurtosis of the DM-SNR curve', ' Skewness of the DM-SNR curve'}
True
```


Tests for q2 (points: None)
```
>>> set(rawdata.columns) == {' Excess kurtosis of the DM-SNR curve', ' Mean of the DM-SNR curve', ' Mean of the integrated profile', ' Skewness of the integrated profile', ' Standard deviation of the integrated profile', 'target_class'}
True
```


Tests for q3 (points: None)
```
>>> N0 == 11375
True
```
```
>>> N1 == 1153
True
```


Tests for q4 (points: None)
```
>>> (Xtrain, Xtest, ytrain, ytest) = train_test_split(rawdata.iloc[:, :-1], rawdata.iloc[:, -1], test_size=0.1, random_state=rng_seed)
```
```
>>> Xtrain.shape == (11275, 5) and Xtest.shape == (1253, 5)
True
```
```
>>> np.allclose(Xtrain.iloc[:5, :5].values, np.array([[114.0625, 49.06544577, 0.26177655, 5.4506689, 5.92887825], [129.015625, 51.37864713, -0.39287668, 4.06688963, 6.99992854], [118.890625, 52.01971297, -0.3249598, 3.32107023, 6.89619648], [130.109375, 50.53699672, -0.1963287, 2.95986622, 8.40783101], [138.8125, 47.61101337, -0.15766978, 2.55267558, 9.81723671]]), 0.01)
True
```
```
>>> np.allclose(Xtest.iloc[:5, :5].values, np.array([[100.710938, 39.5873544, 1.90577125, 3.03595318, 7.74167135], [113.59375, 50.8120676, 0.106782387, 2.64464883, 8.98869263], [76.484375, 40.38351, 1.83528897, 2.72157191, 8.33382076], [116.882812, 52.8183238, -0.392660094, 7.85033445, 4.32790715], [110.078125, 50.8617319, -0.352484187, 1.7173913, 11.1219638]]), 0.01)
True
```


Tests for q5p1 (points: None)
```
>>> np.allclose(np.sort(model.named_steps['scaler'].scale_), [4.52959342, 6.18703268, 6.77911222, 25.66842747, 29.56810699], 0.001)
True
```
```
>>> np.allclose(np.sort(model.named_steps['scaler'].mean_), [1.77042614, 8.33728113, 12.65272831, 46.52346098, 111.00540743], 0.001)
True
```


Tests for q5p2 (points: None)
```
>>> bool(best_params_logreg['model__penalty'] == 'l2' and np.isclose(best_params_logreg['model__C'], 1.2915496650148826, 0.001))
True
```
```
>>> np.allclose(np.sort(result_logreg['l1']['mean_test_accuracy'].values), [0.93073139, 0.95716184, 0.96691783, 0.97082026, 0.97206199, 0.97321497, 0.9745453, 0.9744566, 0.97454527, 0.97454527], 0.001)
True
```
```
>>> np.allclose(np.sort(result_logreg['l1']['mean_test_balanced_accuracy'].values), [0.6219554, 0.76969708, 0.8298979, 0.8559829, 0.86580857, 0.87253708, 0.87979469, 0.87974587, 0.88022896, 0.88022896], 0.001)
True
```
```
>>> np.allclose(np.sort(result_logreg['l2']['mean_test_accuracy'].values), [0.96798216, 0.96833691, 0.97028809, 0.97135244, 0.97179591, 0.97312627, 0.9740131, 0.97454527, 0.97454527, 0.97463397], 0.001)
True
```
```
>>> np.allclose(np.sort(result_logreg['l2']['mean_test_balanced_accuracy'].values), [0.84155829, 0.84397568, 0.85220597, 0.8584528, 0.86392081, 0.87118123, 0.87688771, 0.88022896, 0.88022896, 0.88027778], 0.001)
True
```


Tests for q6 (points: None)
```
>>> np.allclose(np.sort(result_rf['sqrt']['mean_test_accuracy'].values), [0.96993343, 0.97516615, 0.97578697, 0.9764078, 0.97649655], 0.001)
True
```
```
>>> np.allclose(np.sort(result_rf['sqrt']['mean_test_balanced_accuracy'].values), [0.860297, 0.89275702, 0.89745413, 0.8991489, 0.89996863], 0.001)
True
```
```
>>> np.allclose(np.sort(result_rf['log2']['mean_test_accuracy'].values), [0.96993343, 0.97516615, 0.97578697, 0.9764078, 0.97649655], 0.001)
True
```
```
>>> np.allclose(np.sort(result_rf['log2']['mean_test_balanced_accuracy'].values), [0.860297, 0.89275702, 0.89745413, 0.8991489, 0.89996863], 0.001)
True
```
```
>>> best_params_rf == {'model__max_features': 'sqrt', 'model__n_estimators': 75}
True
```


Tests for q7 (points: None)
```
>>> np.allclose(np.sort(result_ab[0.01]['mean_test_accuracy'].values), [0.96869173, 0.96878043, 0.96931251, 0.96993332, 0.97108625], 0.001)
True
```
```
>>> np.allclose(np.sort(result_ab[0.01]['mean_test_balanced_accuracy'].values), [0.85403613, 0.85917348, 0.85926603, 0.85959107, 0.86056515], 0.001)
True
```
```
>>> np.allclose(np.sort(result_ab[0.1]['mean_test_accuracy'].values), [0.97161842, 0.97223932, 0.97268272, 0.97303754, 0.97365839], 0.001)
True
```
```
>>> np.allclose(np.sort(result_ab[0.1]['mean_test_balanced_accuracy'].values), [0.85966167, 0.86430433, 0.86750053, 0.87074414, 0.87321404], 0.001)
True
```
```
>>> best_params_ab == {'model__learning_rate': 0.1, 'model__n_estimators': 100}
True
```


Tests for q8 (points: None)
```
>>> np.allclose(np.sort(result_gbc[0.1]['mean_test_accuracy'].values), [0.97525489, 0.97525494, 0.97569827], 0.001)
True
```
```
>>> np.allclose(np.sort(result_gbc[0.1]['mean_test_balanced_accuracy'].values), [0.89861049, 0.89972225, 0.90102928], 0.001)
True
```
```
>>> np.allclose(np.sort(result_gbc[1.0]['mean_test_accuracy'].values), [0.93002323, 0.93020063, 0.93091028], 0.001)
True
```
```
>>> np.allclose(np.sort(result_gbc[1.0]['mean_test_balanced_accuracy'].values), [0.66417237, 0.66510045, 0.66649721], 0.001)
True
```
```
>>> best_params_gbc == {'model__learning_rate': 0.1, 'model__n_estimators': 100}
True
```


Tests for q9 (points: None)
```
>>> best_name == 'gbc'
True
```
```
>>> bool(np.isclose(test_accuracy, 0.9736632083000798, 0.001))
True
```
```
>>> bool(np.isclose(test_balanced_accuracy, 0.8885775228008237, 0.001))
True
```
