```python
# Initialize Otter
import otter
grader = otter.Notebook("hw5.ipynb")
```

---

<h1><center>SDSE Homework 5 <br><br> Text Classification with Naive Bayes </center></h1>

---

```python
import numpy as np
import matplotlib.pyplot as plt
import pickle
from hashutils import *
```

In this homework we will apply the technique of Naive Bayes classification to the problem of categorizing text-based documents. The dataset is a selection of posts from scikit-learn's ['20 newsgroups' dataset](https://scikit-learn.org/stable/datasets/real_world.html#newsgroups-dataset), which contains some 18000 newsgroup posts in 20 different categories, such as politics, autos, electronics, atheism, and hockey. For simplicity, we will focus on just two of the categories: computer graphics and motorcycles.

```python
with open('hw5_text.pickle','rb') as file:
    Xtrain, ytrain, Xval, yval, categories, vocabulary = pickle.load(file)
```

# 1. Define constants

Define the following variables in terms of quantities loaded from the pickle file. 
+ `N` ... number of documents in the training corpus
+ `K` ... number of document categories
+ `D` ... number of words in the vocabulary

```python
N = ...
K = ...
D = ...
```

```python
grader.check("q1")
```

# 2. Find the number of training documents for each category

Create a Python dictionary called `docs_per_category` that stores the number of documents in the training data under each category.
The keys of `docs_per_category` are the two categories (`comp.graphics` and `rec.motorcycles`). The value for each key is an integer corresponding to the number of documents in that category.

```python
docs_per_category = dict.fromkeys(categories,0)
for category in categories:
    docs_per_category[category] = ...
```

```python
grader.check("q2")
```

# 3. Bag-of-words representation

Our Naive Bayes model will operate on a bag-of-words representation of each document. A bag-of-words representation is simply a set with all of the individual words of the document that are also in the vocabulary. The first task is to write the `to_bow` method, which converts a document into its bag-of-words representation.

The arguments for this method are the document as a string and the vocabulary. It should return a set (`bow`) with the unique words that appear in both the document and the vocabulary. The comments in the method provide steps to follow.

**Hints**
+ The `set` constructor
+ The `add` method on sets. 

```python
def to_bow(doc,vocabulary):
    bow = set()
    
    # Split `doc` at spaces using the the string's `split` method. Obtain a list of strings.
    words_list = ...

    # Keep only unique words from the list, by casting it as a set.
    word_set = ...

    # From that set, store in bow only the ones that are present in the vocabulary.
    ...

    return bow
```

```python
# This code runs 'to_bow' on each of the documents in Xtrain. 
# It produces Xtrain_bow and bag_sizes

# DO NOT MODIFY THIS CODE ...................
Xtrain_bow = np.empty(len(Xtrain),dtype=set)
bag_sizes = np.empty(len(Xtrain),dtype=int)
for i, doc in enumerate(Xtrain):
    bow = to_bow(doc,vocabulary)
    Xtrain_bow[i] = bow
    bag_sizes[i] = len(bow)
# ...........................................
```

```python
grader.check("q3")
```

# 4. Compute the document count for each word in each category
To estimate the parameters of the Naive Bayes model, we will need to know, for each category and each word, the number of documents of the category that contain the word. This is the purpose of the `find_doc_counts_per_word_category` function.

Implement the `find_doc_counts_per_word_category` following the steps provided in the code. This function accepts training data (`Xtrain_bow` and ` ytrain`), as well as the categories and the vocabulary. It produces `doc_counters`, which is a dictionary indexed by category. For each category, `doc_counters[category]` is a dictionary indexed by words in the vocabulary. For each `word` in the vocabulary, `doc_counters[category][word]` is the number of documents of that category that contain that word.

For example, if we run the `find_doc_counts_per_word_category` and save the result as `doccount_per_cat_and_word`, then we can obtain the number of documents in the computer graphics category that contain the word "windows":

```python 
doccount_per_cat_and_word = find_doc_counts_per_word_category(categories,vocabulary,ytrain,Xtrain_bow)
doccount_per_cat_and_word['comp.graphics']['windows']
```


```python
def find_doc_counts_per_word_category(categories,vocabulary,ytrain,Xtrain_bow):

    # Initialize doc_counters
    doc_counters = dict.fromkeys(categories)
    for category in categories:
        doc_counters[category]  = dict.fromkeys(vocabulary,0)

    # Loop through categories.
    for category in categories:

        # Filter Xtrain_bow and keep only the documents of this category
        docs_in_category = ...

        # For each document in this category, increment the doc_counter entry for all vocabulary words found in the document.
        ...

    return doc_counters
```

```python
# Run `find_doc_counts_per_word_category` on every word in the training dataset

# DO NOT MODIFY THIS CODE ...................
doccount_per_cat_and_word = find_doc_counts_per_word_category(categories,vocabulary,ytrain,Xtrain_bow)
# ...........................................
```

```python
grader.check("q4")
```

# 5. Find word frequencies per category

Write the `compute_freq` method. This method takes `doccount_per_cat_and_word`, `ytrain` and the Laplace smoothing factor `alpha` as inputs and computes word and category frequencies.

+ Category frequencies `catfreq[category]`: The category frequency for category $k$ is the proportion of the documents that are of class $k$.

$$\hat{p}_k = \frac{N_k}{N} $$

where $N_k$ is the number of documents in category $k$, and $N$ is the total number of documents. 

+ Word frequencies `wordfreq[category][word]`: The Laplace-smoothed word frequency for a word $d$ and category $k$.

$$\hat{p}_{d,k} = \frac{N_{d,k}+\alpha}{N_k + \alpha D} $$

where $N_{d,k}$ is the number of documents of category $k$ that contain word $d$, and $\alpha$ is the Laplace smoothing factor. 

```python
def compute_freq(doccount_per_cat_and_word,ytrain,categories,vocabulary,alpha):

    K = len(categories)  # number of categories
    D = len(vocabulary)  # number of vocabulary words
    N = len(ytrain)      # number of documents

    # Compute the number of documents in each category. Store it in the dictionary `ndocs`.
    ndocs = dict.fromkeys(categories)
    for category in categories:
        ndocs[category] = ...

    # Compute the category frequenies. For each category, catfreq[category] equals
    # the number of documents of that category (ndocs) divided by the total number of documents.
    catfreq = dict()
    for category, n in ndocs.items():
        catfreq[category] = ...

    # Initialize wordfreq
    wordfreq = dict.fromkeys(categories)
    for category in categories:
        wordfreq[category] = dict.fromkeys(vocabulary)

    # Compute wordfreq
    # For each category ...
    for category in categories:

        # the denominator is the number of documents in that category + alpha*K
        den = ...

        # iterate through items in `doccount_per_cat_and_word` to compute
        # the word frequency for every category and word.
        for word, doccount in doccount_per_cat_and_word[category].items():
            wordfreq[category][word] = ...

    return wordfreq, catfreq
```

```python
# Run `compute_word_log_freq` with alpha=0.01.

# DO NOT MODIFY THIS CODE ...................
wordfreq, catfreq = compute_freq(doccount_per_cat_and_word,ytrain,categories,vocabulary,0.01)
# ...........................................
```

```python
grader.check("q5")
```

# 6. Write the Naive Bayes prediction function.

Use your Naive Bayes model to predict the category of a given test document `doc`. 

Recall that Naive Bayes selects the category as follows:

$$\hat{y} = \underset{k}{\text{argmax}} \hspace{2mm} \log \hat{p}_k +   \sum_{d:x^d=1}  \log \hat{p}_{d,k} +\sum_{d:x^d=0}  \log (1-\hat{p}_{d,k})$$

The arguments of the `predict` method are:
+ `doc`: a single document as a string.
+ `wordfreq`, `catfreq`: the ratios computed in the previous step (with $\alpha=0.01$)
+ `vocabulary`: the vocabulary.

The steps are:
1. Find the BOW representation of `doc`.

2. Loop through categories, for each one compute its score using the above formula.

3. Return the category with the highest score.

```python
def predict(doc, wordfreq, catfreq, vocabulary):

    # 1. Find the BOW representation of doc.
    doc_bow = ...

    # 2. Loop through categories, for each one compute its score, and save it in score_cat.
    score_cat = dict.fromkeys(categories,0)
    for category in categories:
        score = ...
        ...

        score_cat[category] = score

    # 3. Find the category with the highest score.
    maxcat = ...

    return maxcat
```

```python
# Run `predict` on all of the documents in the validation set.

# DO NOT MODIFY THIS CODE ...................
allpred = list()
for x in Xval:
    allpred.append( predict(x, wordfreq, catfreq, vocabulary) )
# ...........................................

allpred[0]
```

```python
grader.check("q6")
```

# 7. Compute the accuracy of the model

The function `compute_accuracy` takes a dataset (`X`,`y`), computes predictions for the documents in `X` using the `predict` function, and computes the accuracy of these predictions with respect to `y`. The accuracy of the model is defined as the number of correct predictions, divided by the total number of predictions. Implement this function. 

```python
def compute_accuracy(X, y, wordfreq, catfreq, vocabulary):

    # count the number of correct predictions
    correct = 0
    for i in ...

    # accuracy is the ratio of correct predictions to total predictions.
    accuracy = ...

    return accuracy
```

```python
compute_accuracy(Xtrain, ytrain, wordfreq, catfreq, vocabulary)
```

```python
grader.check("q7")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> get_hash(N, 2) == '62dca49f0781bf26b4305bddb0414bea'
True
```
```
>>> get_hash(K, 2) == 'd1bd83a33f1a841ab7fda32449746cc4'
True
```
```
>>> get_hash(D, 3) == '6cc3577970e97938a2191e0443243dab'
True
```


Tests for q2 (points: None)
```
>>> get_hash(docs_per_category['comp.graphics'], 3) == '140c93142d2695247f269ee4af3001b0'
True
```
```
>>> get_hash(docs_per_category['rec.motorcycles'], 3) == '73aa23ef8363441b157cde7d6cd0f634'
True
```


Tests for q3 (points: None)
```
>>> get_hash(''.join([''.join(np.sort(list(bow))) for bow in Xtrain_bow])) == '62075cf30caf7a8196c1d8e97ef68d5d'
True
```
```
>>> get_hash(bag_sizes, 3) == '597702cfceb8c3ef0616f259be56a36d'
True
```


Tests for q4 (points: None)
```
>>> doccount_per_cat_and_word = find_doc_counts_per_word_category(categories, vocabulary, ytrain, Xtrain_bow)
>>> len(doccount_per_cat_and_word['rec.motorcycles']) == 548 and doccount_per_cat_and_word['rec.motorcycles']['first'] == 5
True
```
```
>>> doccount_per_cat_and_word = find_doc_counts_per_word_category(categories, vocabulary, ytrain, Xtrain_bow)
>>> testwords = ['according', 'between', 'could', 'explain', 'harley', 'miles', 'source', 'wondering']
>>> get_hash([doccount_per_cat_and_word['comp.graphics'][word] for word in testwords], 2) == '15e5b3e1d64da1dbd9557db56fe89d31'
True
```
```
>>> doccount_per_cat_and_word = find_doc_counts_per_word_category(categories, vocabulary, ytrain, Xtrain_bow)
>>> testwords = ['according', 'between', 'could', 'explain', 'harley', 'miles', 'source', 'wondering']
>>> get_hash([doccount_per_cat_and_word['rec.motorcycles'][word] for word in testwords], 2) == 'e12cdf40fcd60f51698f335252c3ddc4'
True
```


Tests for q5 (points: None)
```
>>> (wordfreq, catfreq) = compute_freq(doccount_per_cat_and_word, ytrain, categories, vocabulary, 0.01)
>>> bool(len(wordfreq) == 2 and len(catfreq) == 2 and np.isclose(wordfreq['comp.graphics']['does'], 0.11522, atol=0.001))
True
```
```
>>> bool(catfreq['comp.graphics'] > 0.3)
True
```
```
>>> get_hash(catfreq['comp.graphics'], 4) == 'afbc48a7ca8d716f9efa7cc993316668'
True
```
```
>>> (wordfreq, catfreq) = compute_freq(doccount_per_cat_and_word, ytrain, categories, vocabulary, 0.01)
>>> testwords = ['according', 'between', 'could', 'explain', 'harley', 'miles', 'source', 'wondering']
>>> get_hash([float(wordfreq['rec.motorcycles'][word]) for word in testwords], 4) == '4c02fc41f94271e20664441d66a3a4b4'
True
```
```
>>> (wordfreq, catfreq) = compute_freq(doccount_per_cat_and_word, ytrain, categories, vocabulary, 0.01)
>>> testwords = ['according', 'between', 'could', 'explain', 'harley', 'miles', 'source', 'wondering']
>>> get_hash([float(wordfreq['comp.graphics'][word]) for word in testwords], 4) == '0282682c12c6023bdb6840f03a69d989'
True
```


Tests for q6 (points: None)
```
>>> allpred = [predict(x, wordfreq, catfreq, vocabulary) for x in Xval]
>>> allpred = ''.join([a[0] for a in allpred])
>>> get_hash(allpred) == '68d7a8dffe1ac1b0fb68d1afc9c45da8'
True
```


Tests for q7 (points: None)
```
>>> acc0 = compute_accuracy(Xtrain, ytrain, wordfreq, catfreq, vocabulary)
>>> bool(acc0 > 0.9 and acc0 < 1.05)
True
```
```
>>> acc1 = compute_accuracy(Xval, yval, wordfreq, catfreq, vocabulary)
>>> get_hash(acc1, 3) == 'a894124cc6d5c5c71afe060d5dde0762'
True
```
